The cation channel mechanisms of subthreshold inward depolarizing currents in the VTA dopaminergic neurons and their roles in the chronic-stress-induced depression-like behavior

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public Review):

Comment 1: One of the only demonstrations of the expression and physiological significance of TRPCs in VTA DA neurons was published by (Rasmus et al., 2011; Klipec et al., 2016) which are not cited in this paper. In their study, TRPC4 expression was detected in a uniformly distributed subset of VTA DA neurons, and TRPC4 KO rats showed decreased VTA DA neuron tonic firing and deficits in cocaine reward and social behaviors. Update: The authors say they have added a discussion of these papers, but I do not see it in the updated manuscript.

We thank the reviewer for the suggestion. The discussion for this has been added (line 557-565).

Comment 2: The authors should report the results (exact data values) of female mice in the results text, or pool the male and female data if the sex differences are not significant.

We agree with reviewer. Some experiments were further redone with female and the data of male and female mice have been reported in the results of text.

Comment 3: The selectivity of drugs should be referred as "selective" rather than "specific".

Thanks, “specific” has been changed to “selective”.

Comment 4: Line 62: typo, "substantia nigra".

Thanks, “substantial nigra” has been changed to “substantia nigra” in line 65.

Comment 5: Line 77: some new studies suggest that NALCN might have voltage dependency

(rectification).

Thanks, description of NALCN voltage dependence has been corrected in line 81-83.

Comment 6: Line 175: change "less" to "fewer".

Thanks, “less” has been changed to “fewer”.

Comment 7: Line 299: choose one - "was not ... or" or "was neither ... nor".

Thanks, this error has been corrected.

Comment 8: In Figure 1Aii and Figure 3Bi, it was not specified in the results text or figure legend that C1-C5 represent individual cell until the legend for Figure 4.

Thanks, these description about gel have been added in the figure legends.

Reviewer #2 (Public Review):

Comment 1: From the previous review, we mentioned that " 'The HCN' as written in line 69 is a bit misleading, as HCN channels in the heart and brain are different members of a family of channels, although as written in the text, it seems that they are identical." This is still the case (now line 73).

We agreed with the reviewer’s comments. The introduction about HCN has been corrected (line 74-78).

Comment 2: The authors state in line 112 that "most of the experiments were also repeated in female mice" - this is true in the case of most electrophysiological experiments, although not behavioral experiments. Authors should amend the statement in line 112 and clarify in the Discussion section which findings are generalizable between sexes; e.g.:

a. Discussion of HCN contribution to VTA DA activity (beginning line 453) should clarify male mice.

b. Similarly, any discussion of behavioral findings should clarify male mice.

We agreed with the reviewer’s comments. The sexes of mice used have been noted in the results and discussion.

Comment 3: The authors' statement in lines 179-183 ("In contrast, fewer GABAergic neuronal markers (Glutamic acid decarboxylase, GAD1/2 and vesicular GABA transporter, VGAT) co-expressed with the DA neurons, which is consistent with previous studies that VTA DA neurons co-expressing GABAergic neuronal markers mainly project to the lateral habenula") is a little confusing - as stated, it seems that the authors are confirming DA/GABA coexpression in VTA-LHb neurons, which is not the case.

We agreed with the reviewer’s comments. We corrected this statement (line 182-186).

Comment 4: Additional information could be included in the Methods section description of Western Blotting procedures - e.g., what thickness of tissue and what size gauge were used to dissect VTA for these experiments?

Thanks. The description of tissue in Western Blotting procedures has been added.

Comment 5:

a. Grammatical errors in line 23 of Abstract (also lines 31-32)

b. "drove" should read "strove" in line 92

c. Grammatical errors in lines 401, 444, and 448

We thank the reviewer for pointing out grammatical errors and we corrected them.

Reviewer #3 (Public Review):

Comment 1: The main strength of this study lies on a comprehensive bottom-up approach ranging from patch-clamp recordings to behavioral tasks. These tasks mainly address anxiety-like behaviors and so-called depression-like behaviors (sucrose choice, forced swim test, tail suspension test). The results gathered by means of these procedures are clearcut. However, the reviewer believes that the authors should be more cautious when interpreting immobility responses to stress (forced swim, tail suspension) as "depression-like" responses. These stress models have been routinely used (and validated) in the past to detect the antidepressant properties of compounds under investigation, which by no means indicates that these are depression models. For readers interested by this debate, I suggest to read e.g. De Kloet and Molendijk (Biol. Pscyhiatry 2021).

We thank the reviewer for the suggestion. We will be more careful and rigorous in the selection of stress models in our subsequent research work.

Editor's note:

Should you choose to revise your manuscript, please include full statistical reporting including exact p-values wherever possible alongside the summary statistics (test statistic and df) and 95% confidence intervals. These should be reported for all key questions and not only when the p-value is less than 0.05.

We have added the full statistical reporting including exact p-values wherever possible alongside the summary statistics (test statistic and df) and 95% confidence intervals into the results and the figure legends of the revised manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The slow-intrinsic-pacemaker dopaminergic (DA) neurons originating in the ventral tegmental area (VTA) is implicated in various mood-and emotion-related disorders, such as anxiety, fear, stress and depression. Abnormal activity of projection-specific VTA DA neurons is the key factor in the development of these disorders. Here, we describe the crucial role for the NALCN and TRPC6, non-selective cation channels in mediating the subthreshold inward depolarizing current and driving the firing of action potentials of VTA DA neurons in physiological condition. Furthermore, we demonstrate that down-regulation of TRPC6 protein expression in the VTA DA neurons likely contributes to the reduced activity of projection-specific VTA DA neuron in chronic mild unpredictable stress (CMUS) depressive mice. In consistent with this, selective knockdown of TRPC6 channels in the VTA DA neurons conferred mice with depression-like behavior. This current study suggests down-regulation of TRPC6 expression/function is involved in reduced VTA DA neuron firing and chronic stress-induced depression-like behavior in the mouse models of chronic stress-induced depression.

Revised version: The slow-intrinsic-pacemaker dopaminergic (DA) neurons originating in the ventral tegmental area (VTA) are implicated in various mood- and emotion-related disorders, such as anxiety, fear, stress and depression. Abnormal activity of projection-specific VTA DA neurons is the key factor in the development of these disorders. Here, we describe the crucial role for the NALCN and TRPC6, non-selective cation channels in mediating the subthreshold inward depolarizing current and driving the firing of action potentials of VTA DA neurons in physiological condition. Furthermore, we demonstrate that down-regulation of TRPC6 protein expression in the VTA DA neurons likely contributes to the reduced activity of projection-specific VTA DA neuron in chronic mild unpredictable stress (CMUS) depressive mice. In consistent with these, selective knockdown of TRPC6 channels in the VTA DA neurons conferred mice with depression-like behavior. This current study suggests down-regulation of TRPC6 expression/function is involved in reduced VTA DA neuron firing and chronic stress-induced depression-like behavior of mice.

Introduction

Dopaminergic (DA) neurons in the ventral tegmental area (VTA) play a key role in mood, reward and emotion-related behaviors[1]. These DA neurons project, through mesocorticolimbic dopaminergic system, to brain regions including the nucleus accumbens (NAc), medial prefrontal cortex (mPFC) and basolateral amygdala (BLA), and modulate these target neuronal circuits through their regulated activity (firing of action potentials)[2, 3]. As such, understandably, the altered functional activity of the VTA DA neurons is found to be the key determinants in abnormal behaviors related to the diseased conditions such as depression- or anxiety-related states[4]. Thus, the study on the mechanism underlying the firing activity of the VTA DA neurons is crucial for understanding the function of the DA circuitry and the pathogenesis of related mental diseases[4–8].

The VTA DA neurons are slow intrinsic pacemakers, which are further modulated by synaptic inputs from multiple excitatory and inhibitory projections[9, 10]. Firing activity controls the release of DA, thus function of the VTA DA neurons. The altered activity of the VTA DA neurons in a pathological state includes altered firing frequency and switch of the firing patterns[7, 11].

Earlier version: The research on the mechanism for spontaneous firing of action potential (AP) is the key to understand functional regulation of the VTA DA neurons. A variety of ion channels (e.g. voltage-gated Na+ channels, high-threshold-activated Ca2+ channels, Kv2, large conductance calcium-activated potassium channels, A-and M-type potassium channels) are reported to be involved in regulation of the spontaneous firing activity of midbrain DA neurons (including the substantial nigra pars compacta, SNc and VTA)[12–22]. However, it is not completely clear how this spontaneous firing is initiated. AP fires when the resting membrane potential (RMP) is depolarized to the activation threshold for the voltage-dependent Na+ channels, which results in opening of Na+ channels and the upstroke of AP[12]. Therefore, a controlled depolarization of the RMP to the activation threshold of AP is the first key step in a cascade leading to firing of AP. In the VTA DA neurons, the identity of ion channels contributing to this subthreshold depolarization has not been unambiguously established.

Revised version: The research on the mechanism for spontaneous firing of action potential (AP) is the key to understand functional regulation of the VTA DA neurons. A variety of ion channels (e.g. voltage-gated Na+ channels, high-threshold-activated Ca2+ channels, Kv2, large conductance calcium-activated potassium channels, A- and M-type potassium channels) are reported to be involved in regulation of the spontaneous firing activity of midbrain DA neurons (including the substantia nigra pars compacta, SNc and VTA)[12–22]. However, it is not completely clear how this spontaneous firing is initiated. AP fires when the resting membrane potential (RMP) is depolarized to the activation threshold for the voltage-dependent Na+ channels, which results in opening of Na+ channels and the upstroke of AP[12]. Therefore, a controlled depolarization of the RMP to the activation threshold of AP is the first key step in a cascade leading to firing of AP. In the VTA DA neurons, the identity of ion channels contributing to this subthreshold depolarization has not been unambiguously established.

Earlier version: In some cells with spontaneous firing of AP, such as cardiomyocytes and certain neurons (including a subset of midbrain SNc DA neurons[23]), the ionic mechanism for the auto-rhythmicity is the HCN, a non-selective cation channel that would depolarize the membrane potential after being activated at a hyperpolarized membrane potential[23]. However, in the VTA DA neurons, results concerning the HCN channels to the spontaneous firing are controversial[15, 23–25].

Revised version: The hyperpolarization-activated cyclic-nucleotide-modulated (HCN1-4) channel gene family, one non-selective cation channels that would be activated at a hyperpolarized membrane potential[23], represent the molecular correlates of the pacemaker current in some cells with spontaneous firing of AP, such as If in cardiomyocytes and Ih in certain neurons including a subset of midbrain SNc DA neurons[23]. However, in the VTA DA neurons, results concerning the HCN channels to the spontaneous firing are controversial[15, 23–25].

Earlier version: Another non-selective cation channel, NALCN, unlike HCN, is a non-voltage dependent channel[26], and is reported to be among the candidates for molecular correlates of the persistent background Na+ leak currents in the VTA DA neurons [27, 28]. The NALCN-mediated background Na+ leak currents are found to be involved in modulation of firing activity in a variety of neurons including DA neurons in the SNc[26, 28–31]. However, it has not been reported if NALCN plays a role in spontaneous firing of the VTA DA neurons, although the VTA DA neurons have larger Na+ leak currents than SNc DA neurons do[15].

Revised version: Another non-selective cation channel, NALCN, unlike HCN, is a voltage-modulated channel[26], and is reported to be among the candidates for molecular correlates of the persistent background Na+ leak currents in the VTA DA neurons [27, 28]. The NALCN-mediated background Na+ leak currents are found to be involved in modulation of firing activity in a variety of neurons including DA neurons in the SNc[28–32]. However, it has not been reported if NALCN plays a role in spontaneous firing of the VTA DA neurons, although the VTA DA neurons have larger Na+ leak currents than SNc DA neurons do[15].

Earlier version: TRP channels are a large class of non-selective cation channels, and are widely distributed in the peripheral and central nervous system. TRPC3, which has a high homology with TRPC6[32], is found to be involved in the regulation of cardiac and SNc DA neuronal excitability, generating depolarizing currents, triggering action potentials and participating in cellular rhythmic firing[31, 33, 34].The role of TRP channels, except TRPC4[35], in the regulation of VTA DA neuronal excitability has not been described.

Revised version: TRP channels are a large class of non-selective cation channels, and are widely distributed in the peripheral and central nervous system. TRPC3, which has a high homology with TRPC6[33], is found to be involved in the regulation of cardiac and SNc DA neuronal excitability, generating depolarizing currents, triggering action potentials and participating in cellular rhythmic firing[32, 34, 35].The role of TRP channels, except TRPC4[36], in the regulation of VTA DA neuronal excitability has not been described.

Earlier version: In this study, we drove to find channel conductance which contribute to the subthreshold depolarization and thus the spontaneous firing of the VTA DA neurons. And furthermore, we investigated if these channels are the molecular mechanism for the altered function activity of the VTA DA neurons related to depression-like behaviors in mouse model of depression. We started by profiling the expression of non-selective cation channels (NSCCs) of the VTA DA neurons using method of Patch-Seq from midbrain DA neurons projecting to different brain regions, and then focused on HCN, NALCN, TRPC6 and TRPV2 channels which are dominantly expressed in these DA neurons.

Revised version: In this study, we strove to find channel conductance which contribute to the subthreshold depolarization and thus the spontaneous firing of the VTA DA neurons. And furthermore, we investigated if these channels are the molecular mechanism for the altered function activity of the VTA DA neurons related to depression-like behaviors in mouse model of depression. We started by profiling the expression of non-selective cation channels (NSCCs) of the VTA DA neurons using method of Patch-Seq from midbrain DA neurons projecting to different brain regions, and then focused on HCN, NALCN, TRPC6 and TRPV2 channels which are dominantly expressed in these DA neurons.

Results

Inflow of extracellular Na+ contribute to the subthreshold depolarization of the VTA DA neurons

Earlier version: One key feature of the midbrain DA neurons is that these neurons are relatively depolarized in the resting condition; the resting membrane potential (RMP) is far from the K+ equilibrium potential (∼-55 mV vs ∼-90 mV)[15], which enable the spontaneous firing. Thus, there must be persistent inward depolarizing conductance counterbalancing the hyperpolarizing K+ conductance which otherwise would maintain the RMP at a hyperpolarized level. The main aim of this study is to identify the subthreshold depolarizing conductances which contribute to the spontaneous firing of the VTA DA neurons. While we initially performed our experiments using male mice, most of the experiments were also repeated using female mice, results of which were reported separately in the following results. We first observed firing frequency and the RMP of the VTA DA neurons using patch clamp recordings. The midbrain DA neurons, marked by dopamine transporter (DAT), mainly located in the VTA and SNc (Fig. 1Ai). We recorded neurons from the VTA brain slices of mice, and performed single-cell PCR post electrophysiologic recordings to identify the DA neurons with established markers for DA neurons (TH, DAT, D2R (type2 dopamine receptor), and GIRK2 (G protein-gated inward rectifier K inwardly rectifying K+ channels)). Although the majority of cells in the VTA area are DA neurons (∼60%), there are still about 30% GABA neurons and a small number of glutamatergic neurons [36, 37]. Glutamatergic-and GABAergic neurons were also single-cell PCR typed by using markers of vGluT2 (vesicular glutamate transporter 2) and GAD1 (glutamic acid decarboxylase 1), respectively (Fig. 1Aii).

Revised version: One key feature of the midbrain DA neurons is that these neurons are relatively depolarized in the resting condition; the resting membrane potential (RMP) is far from the K+ equilibrium potential (∼-55 mV vs ∼-90 mV)[15], which enable the spontaneous firing. Thus, there must be persistent inward depolarizing conductance counterbalancing the hyperpolarizing K+ conductance which otherwise would maintain the RMP at a hyperpolarized level. The main aim of this study is to identify the subthreshold depolarizing conductances which contribute to the spontaneous firing of the VTA DA neurons. While we initially performed our experiments using male mice, most of the electrophysiological experiments and some Western blot experiments were also repeated using female mice, results of which were reported in the following results. We first observed firing frequency and the RMP of the VTA DA neurons using patch clamp recordings. The midbrain DA neurons, marked by dopamine transporter (DAT), mainly located in the VTA and SNc (Fig. 1Ai). We also performed single-cell PCR to identify the DA neurons with established markers for DA neurons (TH, DAT, D2R (type2 dopamine receptor), and GIRK2 (G protein-gated inward rectifier K inwardly rectifying K+ channels)). Although the majority of cells in the VTA area are DA neurons (∼60%), there are still about 30% GABA neurons and a small number of glutamatergic neurons [37, 38]. Glutamatergic- and GABAergic neurons were also single-cell PCR typed by using markers of VGluT2 (vesicular glutamate transporter 2) and GAD1 (glutamic acid decarboxylase 1), respectively (Fig. 1Aii).

Earlier version: The effects of extracellular Ca2+ and Na+ on firing activity of the male VTA DA neurons.

Revised version: The effects of extracellular Ca2+ and Na+ on firing activity of the male and female VTA DA neurons.

Earlier version: A. Identification of the VTA DA neurons. i: confocal images showing the anatomical distribution of DA neurons in the VTA and SNc; the DA neurons were identified to be DAT-immunofluorescence positive (scale bar, 100 µm). ii: left panel: single-cell PCR results; four cells showed the presence of genes indicated. right panel: Percentage of neuronal markers (DA neuron: TH, DAT, D2R, Girk2; Glu neuron: Vglut2 and GABA neuron: GAD1) positive neurons in the VTA neurons (GAPDH positive). n = 63 cells, N = 10. B. The effect of replacing extracellular Ca2+ with Mg2+ on firing frequency of the VTA DA neurons. Example time-course (i) and traces (ii) of the spontaneous firing before and after replacement of extracellular Ca2+ (2 mM) by an equimolar amount of Mg2+.The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). iii: Summarized data for ii. Firing of the VTA DA neurons were recorded using loose cell-attached patch from a brain slice of the VTA (n = 16, N = 8). Paired-sample T test, t = 4.698, P = 0.0003. C. The effect of replacing extracellular Ca2+ with Mg2+ on the resting membrane potential (RMP) of the VTA DA neurons, in the presence of 1 μM TTX. Example time-course (i) and summarized data (ii) for the resting membrane potential before and after replacement of extracellular Ca2+ (2 mM) by an equimolar amount of Mg2+ (n = 14, N = 8). The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Wilcoxon matched-pairs signed rank test, W = 66.00, P = 0.037. D. i: Replacement of external Na+ by equimolar NMDG resulted in hyperpolarization of the resting membrane potential of the VTA DA neurons, in the presence of 1 μM TTX. ii: Summarized data for i (n = 8, N = 5). The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 7.803, P = 0.0001. * P < 0.05, *** P < 0.001. n is number of neurons recorded and N is number of mice used.

Revised version: A. Identification of the VTA DA neurons. i: confocal images showing the anatomical distribution of DA neurons in the VTA and SNc; the DA neurons were identified to be DAT-immunofluorescence positive (scale bar, 100 µm). ii: left panel: single-cell PCR results; four cells (C1-C4) showed the presence of genes indicated. right panel: Percentage of neuronal markers (DA neuron: TH, DAT, D2R, Girk2; Glu neuron: VGluT2 and GABA neuron: GAD1) positive neurons in the VTA neurons (GAPDH positive). n = 102 cells, N = 22. B. The effect of replacing extracellular Ca2+ with Mg2+ on firing frequency of the male and female VTA DA neurons. Example time-course (i) and traces (ii) of the spontaneous firing before and after replacement of extracellular Ca2+ (2 mM) by an equimolar amount of Mg2+.The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). iii: Summarized data for ii. Firing of the VTA DA neurons were recorded using loose cell-attached patch from a brain slice of the VTA (n = 32, N = 16). Wilcoxon matched-pairs signed rank test, W = 526.0, ****P < 0.0001. C. The effect of replacing extracellular Ca2+ with Mg2+ on the resting membrane potential (RMP) of the male and female VTA DA neurons, in the presence of 1 μM TTX. Example time-course (i) and summarized data (ii) for the resting membrane potential before and after replacement of extracellular Ca2+ (2 mM) by an equimolar amount of Mg2+ (n = 25, N = 15). The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 2.368, df = 24, 95% CI: 0.1454 to 2.119, *P = 0.0263. D. Replacement of external Na+ by equimolar NMDG resulted in hyperpolarization of the resting membrane potential of the male and female VTA DA neurons, in the presence of 1 μM TTX. Example time-course (i) and summarized data (ii) for the resting membrane potential before and after replacement of extracellular Na+ (151 mM) with an equimolar amount of NMDG (n = 17, N = 10). The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 7.103, df = 16, 95% CI: −15.35 to −8.293, ****P < 0.0001. * P < 0.05, **** P < 0.0001. n is number of neurons recorded and N is number of mice used.

Earlier version: To focus on the intrinsic channel conductance for spontaneous firing, potential input modulation from fast-type transmitter transmission to the VTA DA neurons, namely activation of AMPA/kainite/NMDA receptor and GABAA receptor were blocked by receptor antagonists CNQX, DL-AP5 and picrotoxin [15], respectively, in the following experiments.

Revised version: To focus on the intrinsic channel conductance for spontaneous firing, potential input modulation from fast-type transmitter transmission to the VTA DA neurons, namely activation of AMPA/kainite/NMDA receptor and GABAA receptor were blocked by receptor antagonists CNQX, APV and picrotoxin [15], respectively, in the following experiments.

Earlier version: We first studied whether inflow of extracellular Ca2+ contributed to initiation of spontaneous firing in the VTA DA neurons. It has been shown that, unlike the SNc DA neurons, the VTA DA neurons do not have subthreshold Ca2+ oscillatory waves[15], suggesting that subthreshold Ca2+ should not be an important component of subthreshold depolarizing conductance in the VTA DA neurons. In line with this finding, we found in our study, replacing Ca2+ with Mg2+ from the extracellular recording solution (ACSF, artificial cerebral spinal solution) did not reduce, but rather instead, increased the firing frequency of the VTA DA neurons in male mice (Fig. 1B, from 1.71 ± 0.19 Hz to 2.36 ± 0.28 Hz, n = 16, N = 8, P < 0.001). Consistent with this, the RMP of the VTA DA neurons was not hyperpolarized by replacing Ca2+ with Mg2+ from ACSF (Fig. 1C, from -50.77 ± 0.85 mV to -49.47 ± 1.09 mV, n = 14, N = 8, P < 0.05).

Revised version: We first studied whether inflow of extracellular Ca2+ contributed to initiation of spontaneous firing in the VTA DA neurons. It has been shown that, unlike the SNc DA neurons, the VTA DA neurons do not have subthreshold Ca2+ oscillatory waves[15], suggesting that subthreshold Ca2+ should not be an important component of subthreshold depolarizing conductance in the VTA DA neurons. In line with this finding, we found in our study, replacing Ca2+ with Mg2+ from the extracellular recording solution (ACSF, artificial cerebral spinal solution) did not reduce, but rather instead, increased the firing frequency of the VTA DA neurons in male and female mice (Fig. 1B, from 1.77 ± 0.14 Hz to 2.24 ± 0.18 Hz, n = 32, N = 16, Wilcoxon matched-pairs signed rank test, W = 526.0, P < 0.0001). Consistent with this, the RMP of the VTA DA neurons was not hyperpolarized by replacing Ca2+ with Mg2+ from ACSF (Fig. 1C, from −50.62 ± 0.62 mV to −49.49 ± 0.87 mV, n = 25, N = 15, Paired-sample T test, t = 2.368, df = 24, 95% CI: 0.1454 to 2.119, P = 0.0263).

Earlier version: We then, in comparison with Ca2+, observed the role of extracellular Na+ in the RMP of the VTA DA neurons. For this, we followed the method of a previous study[15], by replacing Na+ in ACSF with NMDG (N-methyl-d-glucamine); TTX (Tetrodotoxin) was also added to block the TTX-sensitive Na+ currents, which is known to be the main contributor to the suprathreshold depolarization component of the action potential. Replacement of the extracellular Na+ with NMDG resulted in a hyperpolarization of the RMP in mice (from -46.32 ± 1.97 mV to -64.53 ± 1.74 mV, n = 8, N = 5, P < 0.001, Fig. 1D), indicating a persistent inflow of Na+ through a TTX-insensitive conductance which caused significant depolarization of the neurons, contributing significantly to the subthreshold depolarization of the VTA DA neurons. Similar results were obtained in the VTA DA neurons of female mice, with increased firing frequency and hyperpolarized RMP when the extracellular Ca2+and Na+ were replaced, respectively (sFig. 1A and B).

Revised version: We then, in comparison with Ca2+, observed the role of extracellular Na+ in the RMP of the VTA DA neurons. For this, we followed the method of a previous study[15], by replacing Na+ in ACSF with NMDG (N-methyl-d-glucamine); TTX (Tetrodotoxin) was also added to block the TTX-sensitive Na+ currents, which is known to be the main contributor to the suprathreshold depolarization component of the action potential. Replacement of the extracellular Na+ with NMDG resulted in a hyperpolarization of the RMP in male and female mice (from −48.11 ± 0.99 mV to −59.93 ± 1.51mV, n = 17, N = 10, Paired-sample T test, t = 7.103, df = 16, 95% CI: −15.35 to −8.293, P < 0.0001, Fig. 1D), indicating a persistent inflow of Na+ through a TTX-insensitive conductance which caused significant depolarization of the neurons, contributing significantly to the subthreshold depolarization of the VTA DA neurons. Taken together, above results indicate persistent Na+ influx contribute to the subthreshold depolarization of the VTA DA neurons.

Removed in the revised version: Taken together, above results indicate persistent Na+ influx contribute to the subthreshold depolarization of the VTA DA neurons.

Potential candidates of channel conductance contributing to subthreshold depolarization of the VTA DA neurons

Earlier version: The channel conductance mediating the above-described persistent Na+ influx should come from a cation channel(s) permeating to Na+. We thus investigated the presence and expression level of non-selective cation channels (NSCCs) in the VTA DA neurons using a combination of patch clamp and single-cell RNA-Seq (Patch-seq) technology[38]. In this part of the study, we considered the heterogeneity and projection-specificity of the VTA DA neurons, i.e. VTA DA neurons projecting to different brain regions have different electrophysiological characteristics[39]. Using retrograde labelling techniques using retrobeads, sixty male VTA neurons projecting to five different brain regions (medial prefrontal cortex, mPFC; basal lateral amygdala, BLA; nucleus accumbens core, NAc c; nucleus accumbens lateral shell, NAc ls; nucleus accumbens medial shell, NAc ms) (sFig. 2) were collected for RNA-Seq, using patch-clamp electrodes. Consistent with previous study[39], the VTA DA neurons projecting to above different brain regions were anatomically congregated into subregions of the VTA (sFig. 3). The high expression of biomarkers (Tyrosine hydroxylase, TH; Dopamine decarboxylase, Ddc and Dopamine transporter, DAT) in the 45 cells of 60 cells (Fig. 2A) indicates that these are DA neurons. In addition, less than half of sequenced DA neurons (predominantly those projecting to NAc ms, BLA, and mPFC) have biomarkers (Vesicular glutamate transporter, VGluT1-3) of glutamatergic neurons, predicting a possible coexistence of neurotransmitters (DA and glutamate)[40]. In contrast, fewer GABAergic neuronal markers (Glutamic acid decarboxylase, GAD1/2 and vesicular GABA transporter, VGAT) co-expressed with the DA neurons, which is consistent with previous studies that VTA DA neurons co-expressing GABAergic neuronal markers mainly project to the lateral habenula[41]. It needs to be noted that some neurons in the VTA only expressed markers for glutamatergic or GABAergic neurons, these neurons were excluded from further analysis as we focused our study on the DA neurons.

Revised version: The channel conductance mediating the above-described persistent Na+ influx should come from a cation channel(s) permeating to Na+. We thus investigated the presence and expression level of non-selective cation channels (NSCCs) in the VTA DA neurons using a combination of patch clamp and single-cell RNA-Seq (Patch-seq) technology[39]. In this part of the study, we considered the heterogeneity and projection-specificity of the VTA DA neurons, i.e. VTA DA neurons projecting to different brain regions have different electrophysiological characteristics[40]. Using retrograde labelling techniques by retrobeads, sixty male VTA neurons projecting to five different brain regions (medial prefrontal cortex, mPFC; basal lateral amygdala, BLA; nucleus accumbens core, NAc c; nucleus accumbens lateral shell, NAc ls; nucleus accumbens medial shell, NAc ms) (sFig.1) were collected for RNA-Seq, using patch-clamp electrodes. Consistent with previous study[40], the VTA DA neurons projecting to above different brain regions were anatomically congregated into subregions of the VTA (sFig.2). The high expression of biomarkers (Tyrosine hydroxylase, TH; Dopamine decarboxylase, Ddc and Dopamine transporter, DAT) in the 45 cells of 60 cells (Fig. 2A) indicates that these are DA neurons. In addition, fewer than half of sequenced DA neurons (predominantly those projecting to NAc ms, BLA, and mPFC) have biomarkers (Vesicular glutamate transporter, VGluT1-3) of glutamatergic neurons, predicting a possible coexistence of neurotransmitters (DA and glutamate)[41].In contrast, fewer GABAergic neuronal markers (Glutamic acid decarboxylase, GAD1/2 and vesicular GABA transporter, VGAT) co-expressed with the DA neurons, which is consistent with previous studies that VTA DA neurons which co-release GABA, contain low and even undetectable levels of GAD1, GAD2 and VGAT mRNA or protein[42, 43]. It needs to be noted that some neurons in the VTA only expressed markers for glutamatergic or GABAergic neurons, these neurons were excluded from further analysis as we focused our study on the DA neurons. Expression levels of NSCCs in the VTA DA neurons classified by projection specificity were analyzed and shown in Fig. 2B. Different colored strips on the top of the figure corresponded to the projection specificity coded by the same colors in Fig. 2A. Multiple transient receptor potential channels (TRP) were present in the VTA DA neurons, with prominent expression level of TRPC6 and TRPV2 (Fig. 2Bi). Other prominently expressed NSCCs included HCN2, HCN3, NALCN and pannexin 1 (Panx1) (Fig. 2Bii). Summarized averaged relative expression levels for eight most richly expressed NSCCs were shown in Fig. 2Biii. Furthermore, projection-specific expression of these dominantly expressed NSCCs were analyzed and shown in Fig. 2C; interestingly only TRPC6 seemed expressed in a projection-specific manner, namely, the VTA DA neurons projecting to the NAc have higher TRPC6 expression than the VTA DA neurons projecting to the BLA and the mPFC (Fig. 2C). In our subsequent study, we focused our study on HCN, NALCN, TRPC6 and TRPV2, one for they are the prominently expression channels, and two for some of these channels have been indicated in modulation of excitability of different neuron types[23, 24, 28–32, 44]. Panx1 was investigated in a separate study since it normally composes the semi channels which are different from these more conventional ion channels.

Gene expression profile of non-selective cation channels (NSCCs) in the male VTA DA neurons.

Earlier version: Single-cell RNA-Seq was performed on the VTA neurons projecting to five different brain regions of mice (NAc c, NAc ms, NAc ls, BLA, mPFC). A. Gene expression profile of markers for neuron subtypes from 45 VTA neurons; the neuron subtypes included: dopaminergic (Th, Ddc, DAT), GABAergic (Gad1, Gad2, VGAT) and glutamatergic (VGluT2, VGluT1, VGluT3), which was arranged in different rows indicated in the right labels and in the left-colored vertical lines; projection targets of these neurons were indicated at the top by the colored lines and labels. Relative expression levels of these genes were indicated by the dark-blue color intensity which was transformed from the log2 values of the number of transcripts per million (FPKM) plus 1. B. Relative gene expression levels of TRP channels (i) and other NSCCs (ii) in 45 individual VTA DA neurons and the population average (mean, right columns). Each column of the individual neurons in (i) and (ii) corresponded to the columns in A. (iii) Bar graph of the mean log2 (FPKM +1) for the top 8 NSCCs in descending order. Error bars indicate SEM. C. Gene expression profile of aforementioned top 8 NSCCs from 45 VTA neurons.

Revised version: Single-cell RNA-Seq was performed on the VTA neurons projecting to five different brain regions of male mice (NAc c, NAc ms, NAc ls, BLA, mPFC). A. Gene expression profile of markers for neuron subtypes from 45 VTA neurons; the neuron subtypes included: dopaminergic (Th, Ddc, DAT), GABAergic (Gad1, Gad2, VGAT) and glutamatergic (VGluT2, VGluT1, VGluT3), which was arranged in different rows indicated in the right labels and in the left-colored vertical lines; projection targets of these neurons were indicated at the top by the colored lines and labels. Relative expression levels of these genes were indicated by the dark-blue color intensity which was transformed from the log2 values of the number of transcripts per million (FPKM) plus 1. B. Relative gene expression levels of TRP channels (i) and other NSCCs (ii) in 45 individual VTA DA neurons and the population average (mean, right columns). Each column of the individual neurons in (i) and (ii) corresponded to the columns in A. (iii) Bar graph of the mean log2 (FPKM +1) for the top 8 NSCCs in descending order. Error bars indicate SEM. C. Gene expression profile of aforementioned top 8 NSCCs from 45 VTA neurons.

Removed in the revised version: NALCN contributes to subthreshold depolarization and spontaneous firing of the male VTA DA neurons.

Removed in the revised version: A. Confocal images showing co-expression of NALCN (green) and DAT (red) representing DA neurons (scale bar, 10 µm). B. i: Single-cell PCR from the VTA DA neurons with the expression of NALCN. ii: Percentage of NALCN positive neurons from 15 DA neurons (with expression of DAT). C. Example time course (i) and summarized data (ii) showed that NALCN channel blockers (GdCl3) significantly hyperpolarized the resting membrane potential (RMP) (n = 6, N = 6). The inset on the top of Cii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Wilcoxon matched-pairs signed rank test, W = -21.00, P = 0.0313. D. Example time-course and traces (i) and summarized data (ii) of the effect of GdCl3 on spontaneous firing frequency in DA neurons (n = 7, N = 3). The inset on the top of Dii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 5.742, P = 0.0012. E. Example time course (i) and summarized data (ii) showed that NALCN channel blocker (L703,606) significantly hyperpolarized the resting membrane potential (RMP) (n = 6, N = 4). The inset on the top of Eii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 14.95, P = 0.0001. F. Example time-course and traces (i) and summarized data (ii) of the effect of L703,606 on spontaneous firing frequency in DA neurons (n = 8, N = 4). The inset on the top of Fii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 6.236, P = 0.0004. G. shRNA against NALCN carried by AAV virus (i) (NALCN-shRNA) was injected into the VTA of mice, the mRNA level in the shRNA-NALCN transfected VTA (NALCN-KD) and the scramble shRNA transfected VTA (Control) was analyzed using qPCR (ii) (N = 6). Two-sample T test, t = 3.742, P = 0.0038. H. Confocal images showing the expression of AAV9-NALCN-shRNA-GFP (green) in the VTA DA neurons (DAT, red) (scale bar, 10 µm). I. Loose cell-attached current clamp recordings of the spontaneous firing of the VTA DA neurons transfected with either NALCN-shRNA (ii, n = 12, N = 3) or scramble-shRNA (Con, i, n = 10, N = 3) (both GFP and DAT positive). Examples of firing traces in the middle of i and ii and summarized data in the bottom of i were shown. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: NALCN-KD, red dots). Mann-Whitney U test, U = 0, P = 0.0001. J. Whole-cell current clamp recordings of the resting membrane potential (RMP) of the VTA DA neurons transfected with either NALCN-shRNA (NALCN-KD, ii, n = 18, N = 5) or scramble-shRNA (Con, i, n = 15, N = 5) (both GFP and DAT positive). Examples of RMP in the middle of i and ii and summarized data in the bottom of i were shown. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: NALCN-KD, red dots). Mann-Whitney U test, U = 0, P = 0.0001. * P < 0.05, ** P <0.01, *** P < 0.001. n is number of neurons recorded and N is number of mice used.

Removed in the revised version: Expression levels of NSCCs in the VTA DA neurons classified by projection specificity were analyzed and shown in Fig. 2B. Different colored strips on the top of the figure corresponded to the projection specificity coded by the same colors in Fig. 2A. Multiple transient receptor potential channels (TRP) were present in the VTA DA neurons, with prominent expression level of TRPC6 and TRPV2 (Fig. 2Bi). Other prominently expressed NSCCs included HCN2, HCN3, NALCN and pannexin 1 (Panx1) (Fig. 2Bii). Summarized averaged relative expression levels for eight most richly expressed NSCCs were shown in Fig. 2Biii. Furthermore, projection-specific expression of these dominantly expressed NSCCs were analyzed and shown in Fig. 2C; interestingly only TRPC6 seemed expressed in a projection-specific manner, namely, the VTA DA neurons projecting to the NAc have higher TRPC6 expression than the VTA DA neurons projecting to the BLA and the mPFC (Fig. 2C). In our subsequent study, we focused our study on HCN, NALCN, TRPC6 and TRPV2, one for they are the prominently expression channels, and two for some of these channels have been indicated in modulation of excitability of different neuron types[23, 24, 26, 28–31, 42]. Panx1 was investigated in a separate study since it normally composes the semi channels which are different from these more conventional ion channels.

HCN does not contribute to the spontaneous firing of the VTA DA neurons

Earlier version: We mainly used two pharmacological tools as HCN channel blockers, CsCl and ZD7288[15, 24] to explore the role of HCN channels in the spontaneous firing of the male VTA DA neurons. Efficient blocking effect of these blockers on HCN channel was first verified. For this, the sag potential produced by a hyperpolarizing current (-100 pA) injection was used to assess HCN activity[24]. Both CsCl (3 mM) and ZD7288 (60 μM) effectively reduced the sag potential (from 16.43 ± 2.96 mV to 3.84 ± 1.92 mV, n = 4, N = 4, P < 0.05; from 19.38 ± 3.14 mV to 4.18 ± 2.60 mV, n = 4, N = 4 respectively, P < 0.05) (sFig. 4A and B).

Revised version: We mainly used two pharmacological tools as HCN channel blockers, CsCl and ZD7288[15, 24] to explore the role of HCN channels in the spontaneous firing of the male VTA DA neurons. Efficient blocking effect of these blockers on HCN channel was first verified. For this, the sag potential produced by a hyperpolarizing current (−100 pA) injection was used to assess HCN activity[24]. Both CsCl (3 mM) and ZD7288 (60 μM) effectively reduced the sag potential of VTA DA neurons in male mice (from 16.43 ± 2.96 mV to 3.84 ± 1.92 mV, n = 4, N = 4, Paired-sample T test, t = 4.246, df = 3, 95% CI: −22.03 to −3.155, P = 0.0239; from 19.38 ± 3.14 mV to 4.18 ± 2.60 mV, n = 4, N = 4, Paired-sample T test, t = 5.017, df = 3, 95% CI: −24.83 to −5.555, P = 0.0152, respectively) (sFig. 3A and B).

Removed in the revised version: TRP channels contribute to subthreshold depolarization and spontaneous firing of the male VTA DA neurons.

Removed in the revised version: A. The expression of TRP channels in VTA DA neurons. i: Single-cell PCR from 5 VTA cells (C1-C5). ii: Percentage of TRP channels (C3, C6 and V2) positive neurons in the VTA DA neurons (DAT positive). B and C. Example time-course (i), traces (ii) and the summarized data (iii) for the effect of a nonselective cation channel blocker 2-APB (100 μM, n = 12, N = 5) (B) and a potent TRPC6 inhibitor LA (10 μM, n = 7, N = 5) (C) on the spontaneous firing frequency in the VTA DA neurons. The inset on the top of Bii and Cii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). B: Wilcoxon matched-pairs signed rank test, W = -78.00, P = 0.0005. C: Paired-sample T test, t = 4.089, P = 0.0064. D and E. Example time-course (i) and the summarized data (ii) for the effect of a nonselective cation channel blocker 2-APB (n = 6, N = 6) (D) and a potent TRPC6 inhibitor LA (n = 5, N = 4) (E) on the resting membrane potential (RMP) of the VTA DA neurons. Paired-sample T test, D: t = 7.976, P = 0.0005, E: t = 5.610, P = 0.005. ** P < 0.01, *** P < 0.001. n is number of neurons recorded and N is number of mice used.

Removed in the revised version: However, these same HCN blockers did not inhibit the spontaneous firing of the VTA DA neurons; CsCl slightly increased (1.96 ± 0.19 Hz to 2.14 ± 0.34 Hz, n = 7, N = 4) whereas ZD7288 slightly decreased the firing frequency (1.65 ± 0.25 Hz to 1.50 ± 0.24 Hz, n = 13, N = 5), but none of these effects were statistically significant (P > 0.05) (sFig. 4C and D). It has been reported that VTA DA neurons projecting to the NAc lateral shell have more pronounced HCN/Ih currents than other VTA DA neurons[39]. However, even the VTA DA neurons projecting to the NAc lateral shell (retrogradely labelled by retrobeads), no effect of CsCl or ZD7288 on the spontaneous firing of the male VTA DA neurons were observed (from 2.04 ± 0.24 Hz to 1.93 ± 0.22 Hz, n = 5, N = 4, by CsCl, and from 1.71 ± 0.29 Hz to 1.65 ± 0.33 Hz, n = 5, N = 5, by ZD7288, P > 0.05) (sFig. 5A and B).

Earlier version: TRPC6 contributes to subthreshold depolarization and spontaneous firing of the male VTA DA neurons.

Revised version: NALCN contributes to subthreshold depolarization and spontaneous firing of the VTA DA neurons.

Earlier version: A. The normalized expression profile of TRPC6 in NAc c-and mPFC-projecting VTA DA neurons was verified using single cell-qPCR (NAc c: n = 28, N = 7; mPFC: n = 25, N = 8, both TH positive), Mann-Whitney U test, U = 174, P = 0.0014. B. Confocal images showing co-expression of TRPC6 (green) and TH (red) representing DA neurons (scale bar, 10 µm). C. The efficiency of shRNA knockdown of TRPC6 in the VTA verified by qPCR. shRNA against TRPC6 carried by AAV virus (i) (TRPC6-shRNA) was injected into the VTA of mice, the mRNA level in the shRNA-TRPC6 transfected VTA (TRPC6-KD, N = 5) and the scramble shRNA transfected VTA (Con, N = 5) was analyzed using qPCR (ii). Two-sample T test, t=4.184, P=0.0031. D. Immunofluorescence labelling showing the expression of AAV9-shRNA (TRPC6)-GFP (green) and DAT (red) in the VTA (scale bar, 10 µm). E. Loose cell-attached current clamp recordings of the spontaneous firing of the VTA DA neurons from the mice transfected with either TRPC6-shRNA (TRPC6-KD, ii, n = 9, N = 3) or scramble-shRNA (Con, i, n = 9, N = 3) (both GFP and DAT positive). Examples of firing traces in the middle of i and ii and summarized data in the bottom of ii were shown. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: TRPC6-KD, red dots). Two-sample T test, t = 3.005, P = 0.0084. F. Example time-course (i) and summarized data (ii) showed that shRNA knockdown of TRPC6 in the VTA decreased the 2-APB-inhibited firing responses of the VTA DA neurons. The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (TRPC6-KD, red dots). (n = 7, N = 5, both GFP and DAT positive), Wilcoxon matched-pairs signed rank test, W = -18.00, P = 0.1563. G. Whole-cell current clamp recordings of the resting membrane potential (RMP) of the VTA DA neurons transfected with either TRPC6-shRNA (TRPC6-KD, ii, n = 10, N = 6) or scramble-shRNA (Con, i, n = 10, N = 6) (both GFP and DAT positive). Examples of RMP in the middle of i and ii and summarized data in the bottom of i were shown. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: TRPC6-KD, red dots). Mann-Whitney U test, U = 8, P = 0.0007. H. Example time-course (i) and summarized data (ii) showed that shRNA knockdown of TRPC6 in the VTA decreased the LA-inhibited firing responses of the VTA DA neurons. The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (TRPC6-KD, red dots). (n = 8, N = 5, both GFP and DAT positive), Paired-sample T test, t = 0.000, P = 0.9999. n.s. P > 0.05, ** P <0.01, *** P < 0.001. n is number of neurons recorded and N is number of mice used.

Revised version: A. Confocal images showing co-expression of NALCN (green) and DAT (red) representing DA neurons (scale bar, 10 µm). B. i: Single-cell PCR from the VTA DA neurons (C1-C5) with the expression of NALCN. ii: Percentage of NALCN positive neurons from 15 DA neurons (with expression of DAT). C. Example time-course (i) and summarized data (ii) showed that NALCN channel blocker (GdCl3) significantly hyperpolarized the resting membrane potential (RMP) (n = 7, N = 7) of male VTA DA neurons. The inset on the top of Cii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 11.38, df = 6, 95% CI: −24.30 to −15.70, ****P < 0.0001. D. Example time-course and traces (i) and summarized data (ii) of the effect of GdCl3 on spontaneous firing frequency in male VTA DA neurons (n = 7, N = 3). The inset on the top of Dii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 5.711, df = 6, 95% CI: −0.8183 to −0.3274, **P = 0.0012. E. Example time-course (i) and summarized data (ii) showed that NALCN channel blocker (L703,606) significantly hyperpolarized the resting membrane potential (RMP) (n = 6, N = 4) of VTA DA neurons in male mice. The inset on the top of Eii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Paired-sample T test, t = 14.95, df = 5, 95% CI: −5.498 to −3.885, ****P < 0.0001. F. Example time-course and traces (i) and summarized data (ii) of the effect of L703,606 on spontaneous firing frequency in DA neurons of both sexes (n = 20, N = 10). The inset on the top of Fii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Wilcoxon matched-pairs signed rank test, W = −210.0, ****P < 0.0001.G. shRNA against NALCN carried by AAV virus (i) (NALCN-shRNA) was injected into the VTA of mice, the mRNA level in the shRNA-NALCN transfected VTA (NALCN-KD, N = 6) and the scramble shRNA transfected VTA (Control, N = 6) was analyzed using qPCR (ii). Two-sample T test, t = 3.714, df = 10, 95% CI: −1.200 to −0.3000, **P = 0.0040. H. Confocal images showing the expression of AAV9-NALCN-shRNA-GFP (green) in the VTA DA neurons (DAT, red) (scale bar, 10 µm). I. Loose cell-attached current clamp recordings of the spontaneous firing of the male and female VTA DA neurons transfected with either NALCN-shRNA (ii, n = 35, N = 8) or scramble-shRNA (Con, i, n = 30, N = 8) (both GFP and DAT positive). Examples of firing traces in the middle of i and ii and summarized data in the bottom of i were shown. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: NALCN-KD, red dots). Mann-Whitney U test, U = 0, ****P < 0.0001. J. Whole-cell current clamp recordings of the resting membrane potential (RMP) of the male VTA DA neurons transfected with either NALCN-shRNA (NALCN-KD, ii, n = 18, N = 5) or scramble-shRNA (Con, i, n = 15, N = 5) (both GFP and DAT positive). Examples of RMP traces in the middle of i and ii and summarized data in the bottom of i were shown. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: NALCN-KD, red dots). Mann-Whitney U test, U = 0, ****P < 0.0001. * P < 0.05, ** P < 0.01, **** P < 0.0001. n is number of neurons recorded and N is number of mice used.

Added in the revised version: However, these same HCN blockers did not inhibit the spontaneous firing of the VTA DA neurons in male mice; CsCl slightly increased (1.96 ± 0.19 Hz to 2.14 ± 0.34 Hz, n = 7, N = 4, Paired-sample T test, t = 0.7171, df = 6, 95% CI: −0.4325 to 0.7911, P = 0.5003) whereas ZD7288 slightly decreased the firing frequency (1.65 ± 0.25 Hz to 1.50 ± 0.24 Hz, n = 13, N = 5, Paired-sample T test, t = 1.349, df = 12, 95% CI: −0.4033 to 0.09486, P = 0.2022), but none of these effects was statistically significant (P > 0.05) (sFig. 3C and D). It has been reported that VTA DA neurons projecting to the NAc lateral shell have more pronounced HCN/Ih currents than other VTA DA neurons[40]. However, even the male VTA DA neurons projecting to the NAc lateral shell (retrogradely labelled by retrobeads), no effect of CsCl or ZD7288 on the spontaneous firing of the male VTA DA neurons were observed (from 2.04 ± 0.24 Hz to 1.93 ± 0.22 Hz, n = 5, N = 4, Paired-sample T test, t = 0.4798, df = 4, 95% CI: −0.7601 to 0.5361, P = 0.6564 by CsCl, and from 1.71 ± 0.29 Hz to 1.65 ± 0.33 Hz, n = 5, N = 5, Paired-sample T test, t = 0.4421, df = 4, 95% CI:-0.4368 to 0.3168, P = 0.6813 by ZD7288) (sFig. 4A and B).

Added in the revised version: TRP channels contribute to subthreshold depolarization and spontaneous firing of the VTA DA neurons.

Added in the revised version: A. The expression of TRP channels in VTA DA neurons. i: Single-cell PCR from 5 VTA cells (C1-C5). ii: Percentage of TRP channels (C3, C6 and V2) positive neurons in the VTA DA neurons (DAT positive). B and C. Example time-course (i), traces (ii) and the summarized data (iii) for the effect of a nonselective cation channel blocker 2-APB (100 μM, n = 12, N = 5 for male mice) (B) and a potent TRPC6 inhibitor LA (10 μM, n = 14, N = 8 for male and female mice) (C) on the spontaneous firing frequency in the VTA DA neurons. The inset on the top of Bii and Cii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (red dots). Wilcoxon matched-pairs signed rank test, B: W = −78.00, ***P = 0.0005. C: W = −105.0, ***P = 0.0001. D and E. Example time-course (i) and the summarized data (ii) for the effect of a nonselective cation channel blocker 2-APB (n = 10, N = 9) (D) and a potent TRPC6 inhibitor LA (n = 9, N = 7) (E) on the resting membrane potential (RMP) of the male and female VTA DA neurons. Paired-sample T test, D: t = 8.082, df = 9, 95% CI: −11.90 to −6.693, ****P < 0.0001, E: t = 6.431, df = 8, 95% CI: −5.993 to −2.829, ***P = 0.0002. *** P < 0.001, **** P < 0.0001.n is number of neurons recorded and N is number of mice used.

Earlier version: HCN are activated by membrane hyperpolarization, with an activation threshold around -70 ∼ -90 mV (HCN2, HCN3)[43, 44]. At a depolarized RMP like these in the VTA DA neurons (-51.28 ± 1.85 mV, n = 8, N = 8; e.g. -54.50 mV, sFig. 4Ai, and -53.50 mV, sFig. 4Bi), HCN are most likely not activated thus would not participate in generation of spontaneous firing. To further prove this, we tested if a more hyperpolarized RMP would involve HCN in the spontaneous firing of the VTA DA neurons. For this, we first lowered the K+ concentration in the extracellular ACSF from 3 mM to 1.5 mM to increase the gradient between inside and outside cellular K+; this maneuver indeed hyperpolarized the cell membrane from -53.33 ± 2.79 mV to -65.00 ± 3.02 mV (n = 6, N = 6, P < 0.01) (sFig. 5C). Interestingly, under 1.5 mM extracellular K+, the spontaneous firing of the male VTA DA neurons was almost totally blocked by ZD7288 (sFig. 5D), although an initial enhancement was seen in some neurons (sFig. 5Di).

Revised version: HCN are activated by membrane hyperpolarization, with an activation threshold around −70∼-90 mV (HCN2, HCN3)[45, 46]. At a depolarized RMP like these in the VTA DA neurons (−51.28 ± 1.85 mV, n = 8, N=8; e.g. −54.50 mV, sFig. 3Ai, and −53.50mV, sFig. 3Bi), HCN are most likely not activated thus would not participate in generation of spontaneous firing. To further prove this, we tested if a more hyperpolarized RMP would involve HCN in the spontaneous firing of the VTA DA neurons. For this, we first lowered the K+ concentration in the extracellular ACSF from 3 mM to 1.5 mM to increase the gradient between inside and outside cellular K+; this maneuver indeed hyperpolarized the cell membrane from −49.45 ± 0.93 mV to −62.95 ± 2.56 mV (n = 7, N = 7, Paired-sample T test, t = 6.373, df = 6, 95% CI: −18.69 to −8.318, P = 0.0007) (sFig. 4C). Interestingly, under 1.5 mM extracellular K+, the spontaneous firing of the male VTA DA neurons was almost totally blocked by ZD7288 (sFig. 4D), although an initial enhancement was seen in some neurons (sFig. 4Di).

Earlier version: We tested another possible mechanism for HCN’s failure to participate in generation of VTA DA firing. It has been reported that in the midbrain SNc DA neurons, HCN is involved in the spontaneous firing in young but not in adult mice, due to a more hyperpolarization-shifted activation threshold of HCN in SNc DA neurons in adult mice. We tested if this could also be the case in the VTA DA neurons. Indeed, as shown in sFig. 5E, in contrast to what we saw in adult male mice, the spontaneous firing frequency of the VTA DA neurons in young male mice (less than 15 days postnatal) was significantly reduced by ZD7288 (from 1.86 ± 0.29 Hz to 0.98 ± 0.24 Hz, n = 9, N = 5, P < 0.05).

Revised version: We tested another possible mechanism for HCN’s failure to participate in generation of VTA DA firing. It has been reported that in the midbrain SNc DA neurons, HCN is involved in the spontaneous firing in young but not in adult mice, due to a more hyperpolarization-shifted activation threshold of HCN in SNc DA neurons in adult mice. We tested if this could also be the case in the VTA DA neurons. Indeed, as shown in sFig. 4E, in contrast to what we saw in adult male mice, the spontaneous firing frequency of the VTA DA neurons in young male mice (less than 15 days postnatal) was significantly reduced by ZD7288 (from 1.86 ± 0.29 Hz to 0.98 ± 0.24 Hz, n = 9, N = 5, Paired-sample T test, t = 3.345, df = 8, 95% CI: −1.487 to −0.2734, P = 0.0101).

Earlier version: Taken together, above results suggest HCN is not involved in the spontaneous firing of the VTA DA neurons in adult mice, due to a depolarized RMP and a hyperpolarization-shifted activation property.

Revised version: Taken together, above results suggest HCN is not involved in the spontaneous firing of the VTA DA neurons in adult male mice, due to a depolarized RMP and a hyperpolarization-shifted activation property.

NALCN contributes to subthreshold depolarization and spontaneous firing of the VTA DA neurons

The Na+ currents produced by NALCN have been suggested to be an important component of background Na+ currents in multiple central neurons including DA neurons in the substantia nigra[28].We next investigated if NALCN also play a role in setting the RMP and in the spontaneous firing of the VTA DA neurons. Consistent with above RNA-Seq results, both immunofluorescence (Fig. 3A) and single-cell PCR (Fig.3B) results confirmed broad expression of NALCN in the VTA DA neurons.

Earlier version: GdCl3 (100 μM), a nonspecific NALCN blocker [29, 30], hyperpolarized the RMP of the VTA DA neurons from -52.17 ± 4.15 mV to -67.33 ± 2.68 mV (n = 6, N = 6, P < 0.05) (Fig. 3C), reduced the spontaneous firing frequency of the VTA DA neurons from 1.89 ± 0.29 Hz to 1.32 ± 0.37 Hz (n = 7, N = 3, P < 0.01) (Fig. 3D). L703,606, a specific NALCN blocker[31], also hyperpolarized the RMP of the VTA DA neurons from -49.16 ± 1.48 mV to -54.30 ± 1.67 mV (n = 6, N = 4, P < 0.001) (Fig. 3E), reduced the spontaneous firing frequency of the VTA DA neurons from 2.36 ± 0.38 Hz to 0.64 ± 0.32 Hz (n = 8, N = 4, P < 0.001) (Fig. 3F).

Revised version: GdCl3 (100 μM), a non-selective NALCN blocker [30, 31], hyperpolarized the RMP of the male VTA DA neurons from −48.88 ± 2.41 mV to −68.88 ± 1.28 mV (n = 7, N = 7, Paired-sample T test, t = 11.38, df = 6, 95% CI: −24.30 to −15.70, P < 0.0001) (Fig. 3C), reduced the spontaneous firing frequency of the male VTA DA neurons from 1.89 ± 0.29 Hz to 1.32 ± 0.37 Hz (n = 7, N = 3, Paired-sample T test, t = 5.711, df = 6, 95% CI: −0.8183 to −0.3274, P = 0.0012) (Fig. 3D). L703,606 (10 μM), a selective NALCN blocker[32], also hyperpolarized the RMP of the male VTA DA neurons from −49.61 ± 1.48 mV to −54.30 ± 1.67 mV (n = 6, N = 4, Paired-sample T test, t = 14.95, df = 5, 95% CI:-5.498 to −3.885, P < 0.0001) (Fig. 3E), reduced the spontaneous firing frequency of the male and female VTA DA neurons from 1.87 ± 0.23 Hz to 0.34 ± 0.14 Hz (n = 20, N = 10, Wilcoxon matched-pairs signed rank test, W = −210.0, P < 0.0001) (Fig. 3F).

Earlier version: To observe a more specific effect on NALCN, a shRNA against NALCN was used to knockdown NALCN. For this, AAV9 viral construct (AAV9-U6-shRNA (NALCN)-CMV-GFP) was injected into the VTA of a mouse; similar AAV viral construct containing scramble-shRNA of nonsense sequences was utilized as controls. The qPCR results show sufficient knockdown of NALCN mRNA (NALCN-shRNA: 0.32 ± 0.07, N = 6; control scramble-shRNA: 1.07 ± 0.19, N = 6, P < 0.01, Fig. 3G) in the VTA tissue. Immunofluorescence results (Fig. 3H) show efficient infection of the VTA DA neurons (GFP expression in the DAT-positive neurons) by the virus. Knockdown of NALCN in the VTA DA neurons by shRNA almost completely silenced the firing of the male VTA DA neurons (0.07 ± 0.04 Hz vs 1.84 ± 0.25 Hz in shRNA (n = 12, N = 3) and scramble-shRNA (n = 10, N = 3) infected VTA DA neurons, respectively, P < 0.001, Fig. 3I). Furthermore, the RMP of the male VTA DA neurons was significantly hyperpolarized (-60.67 ± 2.33 mV vs -48.9 ± 1.45 mV in shRNA and scramble-shRNA infected VTA DA neurons, respectively, n = 18 and 15, N = 5 and 5, P < 0.001, Fig. 3J). In female mice, L703,606 also reduced the spontaneous firing frequency of the VTA DA neurons and knockdown of NALCN also silenced the firing of the VTA DA neurons (sFig. 6A and B).

Revised version: To observe a more specific effect on NALCN, a shRNA against NALCN was used to knockdown NALCN. For this, AAV9 viral construct (AAV9-U6-shRNA(NALCN)-CMV-GFP) was injected into the VTA of a mouse; similar AAV viral construct containing scramble-shRNA of nonsense sequences was utilized as controls. The qPCR results show sufficient knockdown of NALCN mRNA (NALCN-shRNA: 0.32 ± 0.07, N = 6; control scramble-shRNA: 1.07 ± 0.19, N = 6, Two-sample T test, t = 3.714, df = 10, 95% CI: −1.200 to −0.3000, P = 0.0040, Fig. 3G) in the VTA tissue. Immunofluorescence results (Fig. 3H) show efficient infection of the VTA DA neurons (GFP expression in the DAT-positive neurons) by the virus. Knockdown of NALCN in the VTA DA neurons by shRNA almost completely silenced the firing of the male and female VTA DA neurons (0.08 ± 0.02 Hz vs 1.60 ± 0.15 Hz in NALCN-shRNA (n = 35, N = 8) and scramble-shRNA (n = 30, N = 8) infected VTA DA neurons, respectively, Mann-Whitney U test, U = 0, P < 0.0001, Fig. 3I).Furthermore, the RMP of the male VTA DA neurons was significantly hyperpolarized (−62.13 ± 1.00 mV vs −50.86 ± 0.57 mV in NALCN-shRNA and scramble-shRNA infected VTA DA neurons, respectively, n = 18 and 15, N = 5 and 5, Mann-Whitney U test, U = 0, P < 0.0001, Fig. 3J).

Removed in the revised version: CMUS depression male mice have a decreased firing activity and down-regulated TRPC6 expression in the VTA DA neurons.

Removed in the revised version: A. Experimental procedure timeline. B. i: Sucrose preference test for CMUS mice. Two-sample T test, t = 5.810, P = 0.0001, compared with control mice. ii: Tail suspension test for CMUS mice, Two-sample T test, t = 4.689, P = 0.0005, compared with control mice. C. Loose cell-attached current clamp recordings of the spontaneous firing of the VTA DA neurons from the CMUS (n = 28, N = 5, DAT positive) and the control (n = 18, N = 5, DAT positive) mice. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: CMUS, red dots). Examples of firing traces in the bottom of i as well as in the bottom left of ii and summarized data in the bottom right of ii were shown. Mann-Whitney U test, U = 118, P = 0.0020. D. Representative Western blot assay (i) and summarized data (ii) showing the expression of TRPC6 and GAPDH in the VTA of the control (Con, N = 6) and the CMUS (N = 6) mice. Two-sample T test, t = 4.848, P = 0.0007. E. i and ii: Example traces (up) and time-course (bottom) of firing frequency of mPFC-projecting VTA DA neurons, effect of 2-APB. (i) Con, (ii) CMUS. iii. Summarized effect of 2-APB in the control (n = 12 cells, N = 5, DAT positive) and CMUS (n = 15 cells, N=5, DAT positive) mice. Two-way ANOVA with Sidak’s multiple comparisons test, Con vs. CMUS: F (1, 25) = 3.411, Con-ACSF vs. CMUS-ACSF, P = 0.0001; Con-2-APB vs. CMUS-2-APB, P = 0.0164. ACSF vs. 2-APB: F (1, 25) = 92.43, Con-ACSF vs. Con-2-APB, P = 0.0001; CMUS-ACSF vs. CMUS-2-APB, P = 0.2259. iv. The inhibition rate (%) on firing rate by 2-APB in the control group (n = 12, N=5) and the CMUS group (n = 15, N = 5). Mann-Whitney U test, U = 1, P = 0.0001. The inset on the top of iii and iv shows a map of a coronal midbrain slice indicating the location of mPFC-projecting neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (iii: Con, grey dots; iv: CMUS, red dots). F. i and ii: Example traces (up) and time-course (bottom) of firing frequency of mPFC-projecting VTA DA neurons, effects of LA.(i) Con, (ii) CMUS. iii. Summarized effect of LA on the control (n = 17 cells, N = 8, DAT positive) and CMUS (n = 19 cells, N = 8, DAT positive) mice. Kruslal-Wallis-H test with Dunnett’s multiple comparisons test, H = 53.08, Con-ACSF vs. Con-LA: P = 0.0001, Con-ACSF vs. CMUS-ACSF: P = 0.0023, CON-LA vs. CMUS-LA: P = 0.0187, CMUS-ACSF vs. CMUS-LA: P > 0.9999. iv. The inhibition rate (%) on firing rate by LA in the control group (n = 17, N = 8) and the CMUS group (n = 19, N = 8). Mann-Whitney U test, U = 0, P = 0.0001. The inset on the top of iii and iv shows a map of a coronal midbrain slice indicating the location of mPFC-projecting neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (iii: Con, grey dots; iv: CMUS, red dots). n.s. P > 0.05, * P < 0.05, **P < 0.01, *** P < 0.001. n is number of neurons recorded and N is the number of mice used.

Taken together, above results suggest that NALCN is a major contributor to the subthreshold depolarization and contribute significantly to the generation of spontaneous firing in the VTA DA neurons.

TRPC6 contributes to subthreshold depolarization and spontaneous firing of the VTA DA neurons

Above RNA-Seq results suggest a broad and strong expression of TRPC6 and TRPV2 channels in the VTA DA neurons. This part of the experiments was focused on the role of these two TRP channels in subthreshold depolarization and spontaneous firing of the VTA DA neurons.

Earlier version: Consistent with the RNA-Seq results, single-cell PCR experiments confirmed high proportion expression of TRPC6 and TRPV2 in the VTA DA neurons (Fig. 4A); among the 28 VTA DA neurons (DAT+), 20 neurons expressed TRPV2 (71.4%), and 18 neurons (64.3%) expressed TRPC6. TRPC3, which is highly homologous with TRPC6[45] and has been reported to be involved in the firing activity of the SNc DA neurons[31], was not detected in RNA-Seq study (Fig. 2B), nor in the single-cell PCR experiment (Fig. 4A). These results reciprocally confirmed the reliability of both RNA-Seq and single cell PCR results.

Revised version: Consistent with the RNA-Seq results, single-cell PCR experiments confirmed high proportion expression of TRPC6 and TRPV2 in the VTA DA neurons (Fig. 4A); among the 28 VTA DA neurons (DAT+), 20 neurons expressed TRPV2 (71.4%), and 18 neurons (64.3%) expressed TRPC6. TRPC3, which is highly homologous with TRPC6[47] and has been reported to be involved in the firing activity of the SNc DA neurons[32], was neither detected in RNA-Seq study (Fig. 2B), nor in the single-cell PCR experiment (Fig. 4A). These results reciprocally confirmed the reliability of both RNA-Seq and single cell PCR results.

Earlier version: General roles of TRP channels in subthreshold depolarization and spontaneous firing of the VTA DA neurons were first assessed using a non-specific broad-spectrum TRP channel blocker, 2-aminoethoxydiphenylborane (2-APB; 100 μM) and a specific TRPC6 channel blocker, larixyl acetate[46] (LA; 10 μM). Both blockers significantly reduced the firing frequency of the male VTA DA neurons (2-APB: from 1.32 ± 0.27 Hz to 0.06 ± 0.02 Hz, n = 12, N = 5, P < 0.001; LA: from 1.16 ± 0.21 Hz to 0.39 ± 0.09 Hz, n = 7, N = 5, P < 0.01; Fig. 4B, 4C). Accordingly, both blockers also significantly hyperpolarized the RMP of the male VTA DA neurons (2-APB: from -51.01 ± 2.20 mV to -61.64 ± 1.28 mV, n = 6, N = 6, P < 0.001; LA: from -53.93 ± 1.05 mV to -57.48 ± 1.09 mV, n = 5, N = 4, P < 0.01; Fig. 4D, 4E).

Revised version: General roles of TRP channels in subthreshold depolarization and spontaneous firing of the VTA DA neurons were first assessed using a non-selective broad-spectrum TRP channel blocker, 2-aminoethoxydiphenylborane (2-APB; 100 μM) and a selective TRPC6 channel blocker, larixyl acetate[48] (LA; 10 μM). Both blockers significantly reduced the firing frequency of the VTA DA neurons (2-APB for male mice: from 1.32 ± 0.27 Hz to 0.06 ± 0.02 Hz, n = 12, N = 5, Wilcoxon matched-pairs signed rank test, W = −78.0, P = 0.0005; LA for male and female mice: from 1.59 ± 0.26 Hz to 0.55 ± 0.15Hz, n = 14, N = 8, Wilcoxon matched-pairs signed rank test, W = −105.0, P = 0.0001; Fig. 4B, 4C). Accordingly, both blockers also significantly hyperpolarized the RMP of the male and female VTA DA neurons (2-APB: from −49.97 ± 1.29 mV to −59.27 ± 1.22 mV, n = 10, N = 9, Paired-sample T test, t = 8.082, df = 9, 95% CI: −11.90 to −6.693, P < 0.0001; LA: from −53.41 ± 0.81 mV to −57.82 ± 1.13 mV, n = 9, N = 7, Paired-sample T test, t = 6.431, df = 8, 95% CI: −5.993 to −2.829, P = 0.0002; Fig. 4D, 4E).

Earlier version: We then tested a more TRPV channel-selective blocker ruthenium red (RR), to study possible role of TRPV2 (and other TRPV) in the spontaneous firing of the VTA DA neurons. RR (60 μM), on average statistically, did not affect the spontaneous firing of the male VTA DA neurons (from 2.30 ± 0.31 Hz to 2.18 ± 0.32 Hz, n = 9, N = 6, P > 0.05) (sFig. 7). However, it needs to be noted that among the nine VTA DA neurons we tested, both increase (sFig. 7B) and decrease (sFig. 7C) of firing frequency by RR were seen, thus with opposite changes of firing frequency in different populations of neurons (sFig. 7A), no overall effect of RR was obtained from the current analysis.

Revised version: We then tested a more TRPV channel-selective blocker ruthenium red (RR), to study possible role of TRPV2 (and other TRPV) in the spontaneous firing of the VTA DA neurons. RR (60 μM), on average statistically, did not affect the spontaneous firing of the male VTA DA neurons (from 2.30 ± 0.31 Hz to 2.18 ± 0.31 Hz, n = 9, N = 6, Paired-sample T test, t = 0.4486, df = 8, 95% CI: −0.7621 to 0.5139, P = 0.6656) (sFig. 5). However, it needs to be noted that among the nine VTA DA neurons we tested, both increase (sFig. 5B) and decrease (sFig. 5C) of firing frequency by RR were seen, thus with opposite changes of firing frequency in different populations of neurons (sFig. 5A), no overall effect of RR was obtained from the current analysis.

Added in the revised version: TRPC6 contributes to subthreshold depolarization and spontaneous firing of the VTA DA neurons.

Added in the revised version: A. The normalized expression profile of TRPC6 in NAc c- and mPFC-projecting VTA DA neurons was verified using single cell-qPCR (NAc c: n = 28, N = 7; mPFC: n = 25, N=8, both TH positive), Mann-Whitney U test, U = 173, **P = 0.0013. B. Confocal images showing co-expression of TRPC6 (green) and TH (red) representing DA neurons (scale bar, 10 µm). C. The efficiency of shRNA knockdown of TRPC6 in the VTA verified by qPCR. shRNA against TRPC6 carried by AAV virus (i) (TRPC6-shRNA) was injected into the VTA of mice, the mRNA level in the shRNA-TRPC6 transfected VTA (TRPC6-KD, N = 5) and the scramble shRNA transfected VTA (Con, N = 5) was analyzed using qPCR (ii). Two-sample T test, t = 4.223, df = 8, 95% CI: −0.8194 to −0.2406, **P = 0.0029. D. Immunofluorescence labelling showing the expression of AAV9-shRNA(TRPC6)-GFP (green) and DAT (red) in the VTA (scale bar, 10 µm). E. Loose cell-attached current clamp recordings of the spontaneous firing of the VTA DA neurons from the mice of both sexes transfected with either TRPC6-shRNA (TRPC6-KD, ii, n = 37, N = 12) or scramble-shRNA (Con, i, n = 27, N = 11) (both GFP and DAT positive). Examples of firing traces in the middle of i and ii and summarized data in the bottom of i were shown. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: TRPC6-KD, red dots). Mann-Whitney U test, U = 48, ****P < 0.0001. F. Example time-course (i) and summarized data (ii) showed that shRNA knockdown of TRPC6 in the VTA of male mice decreased the 2-APB-inhibited firing responses of the VTA DA neurons. The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (TRPC6-KD, red dots). (n = 7, N = 5, both GFP and DAT positive), Wilcoxon matched-pairs signed rank test, W = −18.00, P = 0.1563. G. Whole-cell current clamp recordings of the resting membrane potential (RMP) of the male mice VTA DA neurons transfected with either TRPC6-shRNA (TRPC6-KD, ii, n = 10, N = 6) or scramble-shRNA (Con, i, n = 10, N = 6) (both GFP and DAT positive). Examples of RMP traces in the middle of i and ii and summarized data in the bottom of i were shown. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: TRPC6-KD, red dots). Mann-Whitney U test, U = 8, ***P = 0.0007. H. Example time-course (i) and summarized data (ii) showed that shRNA knockdown of TRPC6 in the male mice VTA decreased the LA-inhibited firing responses of the VTA DA neurons. The inset on the top of ii shows a map of a coronal midbrain slice indicating the location of GFP+ neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (TRPC6-KD, red dots). (n = 8, N = 5, both GFP and DAT positive), Paired-sample T test, t = 0.000, df = 7, 95% CI: −0.05740 to 0.05740, P > 0.9999. n.s. P > 0.05, ** P < 0.01, *** P < 0.001, ****P < 0.0001. n is number of neurons recorded and N is number of mice used.

Added in the revised version: We then focused our next study on TRPC6. Consistent with both RNA-Seq (Fig. 2B) and single-cell PCR (Fig. 4A) results, the mRNA level of TRPC6 in mPFC-projecting VTA single DA neuron is significantly lower than that in NAc c-projecting VTA single DA neuron (0.56 ± 0.12, n = 25, N = 8 for mPFC-projecting TH-positive cells vs 1.09 ± 0.10, n = 28, N = 7 for NAc c-projecting TH-positive cells, Mann-Whitney U test, U = 173, P = 0.0013, Fig. 5A) The immunofluorescence results also demonstrated strong expression of TRPC6 protein in the VTA DA neurons (Fig. 5B, TH-positive). A AAV viral construct with shRNA against TRPC6 (AAV9-U6-shRNA(TRPC6)-CMV-GFP) was injected into the VTA of mice, to knockdown TRPC6; efficiency of this knockdown was assessed by using qPCR measuring the mRNA level of TRPC6 in the VTA tissue (Fig. 5C), and efficiency of viral infection into the VTA DA neurons (marked by DAT) was observed through visualization of GFP (Fig. 5D). Both qPCR (Fig.5C) and immunofluorescence (Fig. 5D) results indicated a sufficient repression of TRPC6 in the VTA DA neurons. Indicative of a significant role in the spontaneous firing and subthreshold depolarization, knockdown of TRPC6 (TRPC6-KD) resulted in a substantial reduction in the spontaneous firing frequency (0.43 ± 0.07 Hz, n = 37, N = 12 for TRPC6-KD vs 1.88 ± 0.16 Hz, n = 27, N = 11 for scramble-shRNA, Mann-Whitney U test, U = 48, P < 0.0001,for male and female mice, Fig. 5E), and hyperpolarization of the RMP (−58.32 ± 1.77 mV, n = 10, N = 6 for TRPC6-KD vs −50.24 ± 0.65 mV, n = 10, N = 6 scramble-shRNA, Mann-Whitney U test, U = 8, P = 0.0007, for male mice, Fig. 5G) of the VTA DA neurons. In addition, the inhibitory effect of 2-APB and LA on the firing of VTA DA neurons infected with TRPC6-shRNA largely diminished, with no statistical difference in firing frequency before and after dosing (2-APB: from 0.58 ± 0.15 Hz to 0.40 ± 0.15 Hz, n = 7, N = 5, Wilcoxon matched-pairs signed rank test, W = −18.00, P = 0.1563, for male mice, Fig. 5F; LA: from 0.43 ± 0.05 Hz to 0.43 ± 0.06 Hz, n = 8, N = 5, Paired-sample T test, t = 0.000, df = 7, 95% CI:-0.05740 to 0.05740, P > 0.9999, for male mice, Fig. 5H).

Added in the revised version: Consequence of knocking down TRPC6 described above should not be a result of a secondary effect on other TRPC channels with which TRPC6 is known to interact to form heteromers, such as TRPC4 and TRPC7[49, 50], since in experiments using single cell PCR (sFig. 6A), it was found only a very small proportion of TRPC6-positive DA cells (TRPC6+DAT+) expressed TRPC4 (sFig. 6Bi) or TRPC7 (sFig. 6Bii), in consistent with the results of single-cell RNA-seq results (Fig. 2). Taken together, above results indicate TRPC6 is an important contributor in subthreshold depolarization and spontaneous firing of the VTA DA neurons.

Added in the revised version: CMUS depression mice have a decreased firing activity and down-regulated TRPC6 expression in the VTA DA neurons.

Added in the revised version: A. Experimental procedure timeline. B. i: Sucrose preference test for CMUS male mice (Con: N = 10; CMUS: N = 10). Two-sample T test, t = 3.569, df = 18, 95% CI: −20.73 to −5.369, **P = 0.0022, compared with control mice. ii: Tail suspension test for CMUS male mice (Con: N = 10; CMUS: N = 10), Two-sample T test, t = 5.826, df = 18, 95% CI: 23.21 to 49.39, ****P < 0.0001, compared with control mice. C. Loose cell-attached current clamp recordings of the spontaneous firing of the VTA DA neurons from the CMUS (n = 68, N = 33, DAT positive) and the control (n = 53, N = 31, DAT positive) mice of both sexes. The inset on the top of i and ii shows a map of a coronal midbrain slice indicating the location of neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (i: Con, grey dots; ii: CMUS, red dots). Examples of firing traces in the bottom of i as well as in the bottom left of ii and summarized data in the bottom right of ii were shown. Mann-Whitney U test, U = 415, ****P < 0.0001. D. Representative Western blot assay (i) and summarized data (ii) showing the expression of TRPC6 and GAPDH in the VTA of the control (Con, N = 10) and the CMUS (N = 11) mice of both sexes. Two-sample T test, t = 5.134, df = 19, 95% CI: −0.7480 to −0.3147, ****P < 0.0001. E. i and ii: Example traces (up) and time-course (bottom) of firing frequency of mPFC-projecting VTA DA neurons, effect of 2-APB. (i) Con and (ii) CMUS of both sexes. iii. Summarized effect of 2-APB in the control (n = 22 cells, N = 15, DAT positive) and CMUS (n = 25 cells, N = 15, DAT positive) mice. Kruskal-Wallis-H test with Dunnett’s multiple comparisons test, Kruskal-Wallis statistic = 66.03, Con-ACSF vs. CMUS-ACSF, **P = 0.0056; Con-2-APB vs. CMUS-2-APB, **P = 0.0016; Con-ACSF vs. Con-2-APB, ****P < 0.0001; CMUS-ACSF vs. CMUS-2-APB, P > 0.9999. iv. The inhibition rate (%) on firing rate by 2-APB in the control group (n = 22, N = 15) and the CMUS group (n = 25, N = 15). Mann-Whitney U test, U = 7, ****P < 0.0001. The inset on the top of iii and iv shows a map of a coronal midbrain slice indicating the location of mPFC-projecting neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (iii: Con, grey dots; iv: CMUS, red dots). F. i and ii: Example traces (up) and time-course (bottom) of firing frequency of mPFC-projecting VTA DA neurons, effects of LA. (i) Con and (ii) CMUS of both sexes. iii. Summarized effect of LA on the control (n = 25 cells, N = 16, DAT positive) and CMUS (n = 30 cells, N = 18, DAT positive) mice. Kruskal-Wallis-H test with Dunnett’s multiple comparisons test, Kruskal-Wallis statistic = 79.59, Con-ACSF vs. CMUS-ACSF: ****P < 0.0001, CON-LA vs. CMUS-LA: **P = 0.0026, Con-ACSF vs. Con-LA: ****P < 0.0001, CMUS-ACSF vs. CMUS-LA: P > 0.9999. iv. The inhibition rate (%) on firing rate by LA in the control group (n = 25, N = 16) and the CMUS group (n = 30, N = 18). Mann-Whitney U test, U = 2, ****P < 0.0001. The inset on the top of iii and iv shows a map of a coronal midbrain slice indicating the location of mPFC-projecting neurons that were recorded and subsequently identified as DA neurons which were DAT positive with single cell-PCR (iii: Con, grey dots; iv: CMUS, red dots). n.s. P > 0.05, **P < 0.01, ****P < 0.0001.n is number of neurons recorded and N is the number of mice used.

Added in the revised version: Down-regulation of TRPC6 contributes to the altered firing activity of the VTA DA neurons in depression model

Added in the revised version: In multiple depression models, the depression-like behavior was directly linked to the altered firing activity of the VTA DA neurons[4, 7, 51–53]. In consideration of the evidence we described above that NALCN and TRPC6 play key roles in firing activity of the VTA DA neurons, we went further to study if these channels also contributed to the altered firing activity of the VTA DA neurons and to the development of the depression-like behavior in depression models of mice.

Added in the revised version: We first established a mice depression model of chronic mild unpredictable stress (CMUS) [54], which, after 5 weeks subjecting to two different stressors every day (Fig. 6A), manifested chronic-stress-induced depression-like behaviors in the sucrose preference test (SPT) and the tail suspension test (TST), with reduced sucrose preference and lengthened immobility time, respectively (male mice: Fig. 6B, female mice: sFig. 8B, I). Multiple other behaviors tests were also performed on these CMUS male and female mice (male mice: sFig. 7, female mice: sFig. 8A, C-H), all indicating depression/anxiety-like behaviors.

Selective knockdown of TRPC6 in the VTA DA neurons confers the male mice with depression-like and anxiety-like behaviors which are reversed by overexpression of TRPC6.

Earlier version: A. shRNA against TRPC6 carried by AAV-loxp virus (TRPC6-shRNA) was injected into the VTA of the DAT-Cre mice. B. The mRNA level in the shRNA-TRPC6 transfected VTA (TRPC6-cKD) and the scramble shRNA transfected VTA (Con) was analyzed using qPCR (Con: N = 4; TRPC6-cKD: N = 6). Mann-Whitney U test, U = 0, P = 0.0159. C. Immunofluorescence labelling showing the expression of AAV9-hSyn-DIO-shRNA (TRPC6)-RFP (red) and DAT (green) in the VTA of the DAT-Cre mice (scale bar, 10 µm). D. The effects of Cre-induced conditional knockdown of TRPC6 in the VTA DA neurons on the behaviors of mice in the sucrose preference test (i), the tail suspension test (ii), and the elevated plus-maze test (iii, iv) (Con: N = 9; TRPC6-cKD: N = 13). Two-sample T test, Sucrose preference test: t = 4.663, P = 0.0001; Tail suspension test: t = 7.143, P = 0.0001; Elevated plus-maze test: open arm t = 10.36, P = 0.0001, closed arm t = 11.18, P = 0.0001.E. The effects of TRPC6 over-expression. AAV9-DIO-TRPC6-GFP (TRPC6-cOE) was injected into the VTA of the DAT-Cre mice, 7 days after TRPC6 knocking down (AAV9-hSyn-DIO-shRNA (TRPC6)-RFP injection). F. The protein level TRPC6 of VTA (Con: the scramble shRNA for TRPC6, N = 7; C6-cKO: TRPC6-cKD, N = 8; c6-cKO+cOE: TRPC6-cKD+ TRPC6-cOE, N = 7). Kruskal-Wallis H test with Dunnett’s multiple comparisons test, H = 14.75, Con vs. C6-cKO: P = 0.0016, C6-cKO vs. C6-cKO+cOE: P = 0.0062. G. Immunofluorescence labelling showing the expression of AAV9-hSyn-DIO-shRNA (TRPC6)-RFP (red), AAV9-DIO-TRPC6-GFP (green) and DAT (blue) in the VTA of the DAT-Cre mice (scale bar, 10 µm). H. The effects of TRPC6 over-expression on depression-like behaviors of TRPC6-cKD mice: the sucrose preference test (i), the tail suspension test (ii), and the elevated plus-maze test (iii) (Con: N = 11; TRPC6-cKD: N = 11; TRPC6-cKD+TRPC6-cOE: N = 11). One-way ANOVA with Dunnett’s multiple comparisons test, sucrose preference test: F (2, 30) = 37.22, Con vs. C6-cKO: P = 0.0001, C6-cKO vs. C6-cKO+cOE: P = 0.0001; tail suspension test: F (2, 30) = 20.10, Con vs. C6-cKO: P = 0.0146, C6-cKO vs. C6-cKO+cOE: P = 0.0001; elevated plus-maze test: F (2, 30) = 17.10, Con vs. C6-cKO: P = 0.0001, C6-cKO vs. C6-cKO+cOE: P = 0.0095. * P < 0.05, ** P < 0.01, *** P < 0.001.n is number of neurons recorded and N is the number of mice used.

Revised version: A. shRNA against TRPC6 carried by AAV-loxp virus (TRPC6-shRNA) was injected into the VTA of the DAT-Cre male mice. B. The mRNA level in the shRNA-TRPC6 transfected VTA (TRPC6-cKD) and the scramble shRNA transfected VTA (Con) was analyzed using qPCR (Con: N = 8; TRPC6-cKD: N = 6). Mann-Whitney U test, U = 0, ***P = 0.0003. C. Immunofluorescence labelling showing the expression of AAV9-hSyn-DIO-shRNA(TRPC6)-RFP (red) and DAT (green) in the VTA of the DAT-Cre mice (scale bar, 10 µm). D. The effects of Cre-induced conditional knockdown of TRPC6 in the VTA DA neurons on the behaviors of mice in the sucrose preference test (i), the tail suspension test (ii), and the elevated plus-maze test (iii, iv) (Con: N = 10; TRPC6-cKD: N = 10). Sucrose preference test: Two-sample T test, t = 4.172, df = 18, 95% CI: −17.07 to −5.634, ***P = 0.0006; Tail suspension test: Two-sample T test, t = 3.377, df = 18, 95% CI: 7.785 to 33.41, **P = 0.0034; Elevated plus-maze test: Mann-Whitney U test, open arm, U = 2, ****P < 0.0001, closed arm, U = 13, **P = 0.0039. E. The effects of TRPC6 over-expression in the TRPC6-cKD-induced depression-like behavior. AAV9-DIO-TRPC6-GFP (TRPC6-cOE) was injected into the VTA of the male DAT-Cre mice, 7 days after injection of AAV9-hSyn-DIO-shRNA(TRPC6)-RFP (TRPC6-cKD). F. The protein level TRPC6 of VTA (Con: the scramble shRNA for TRPC6, N = 7; C6-cKD: TRPC6-cKD, N = 10; C6-cKD+cOE: TRPC6-cKD+TRPC6-cOE, N = 7). Kruskal-Wallis H test with Dunnett’s multiple comparisons test, Kruskal-Wallis statistic = 16.92, Con vs. C6-cKD: ***P = 0.0009, C6-cKD vs. C6-cKD+cOE: **P = 0.0034. G. Immunofluorescence labelling showing the expression of AAV9-hSyn-DIO-shRNA(TRPC6)-RFP (red), AAV9-DIO-TRPC6-GFP (green) and DAT (blue) in the VTA of the DAT-Cre mice (scale bar, 10 µm). H. The effects of TRPC6 over-expression on depression-like behaviors of TRPC6-cKD male mice: the sucrose preference test (i), the tail suspension test (ii), and the elevated plus-maze test (iii) (Con: N = 10; C6-cKD: N = 10; C6-cKD+ cOE: N= 10). the sucrose preference test: Kruskal-Wallis H test with Dunnett’s multiple comparisons test, Kruskal-Wallis statistic = 12.96, Con vs. C6-cKD: **P = 0.0025, C6-cKD vs. C6-cKD+cOE: *P = 0.0139; tail suspension test: One-way ANOVA with Dunnett’s multiple comparisons test, F (DFn, DFd) = 0.8336(2, 27), Con vs. C6-cKD: **P = 0.0011, C6-cKD vs. C6-cKD+cOE: ***P = 0.0003; elevated plus-maze test: One-way ANOVA with Dunnett’s multiple comparisons test, F(DFn, DFd) = 4.884(2, 27), Con vs. C6-cKD: ****P < 0.0001, C6-cKD vs. C6-cKD+cOE: **P = 0.0023. *P < 0.05, **P < 0.01, ***P < 0.001, ****P < 0.0001.n is number of neurons recorded and N is the number of mice used.

Removed in the revised version: We then focused our next study on TRPC6. Consistent with both RNA-Seq (Fig. 2B) and single-cell PCR (Fig. 4A) results, the mRNA level of TRPC6 in mPFC-projecting VTA single DA neuron is significantly lower than that in NAc c-projecting VTA single DA neuron (0.56 ± 0.12, n = 25, N = 8 for mPFC-projecting TH-positive cells vs 1.09 ± 0.10, n = 28, N = 7 for NAc c-projecting TH-positive cells,** P < 0.01, Fig. 5A) The immunofluorescence results also demonstrated strong expression of TRPC6 protein in the VTA DA neurons (Fig. 5B, TH-positive). A AAV viral construct with shRNA against TRPC6 (AAV9-U6-shRNA (TRPC6)-CMV-GFP) was injected into the VTA of mice, to knockdown TRPC6; efficiency of this knockdown was assessed by using qPCR measuring the mRNA level of TRPC6 in the VTA tissue (Fig. 5C), and efficiency of viral infection into the VTA DA neurons (marked by DAT) was observed through visualization of GFP (Fig. 5D). Both qPCR (Fig. 5C) and immunofluorescence (Fig. 5D) results indicated a sufficient repression of TRPC6 in the VTA DA neurons. Indicative of a significant role in the spontaneous firing and subthreshold depolarization, knockdown of TRPC6 (TRPC6-KD) resulted in a substantial reduction in the spontaneous firing frequency (0.80 ± 0.25 Hz, n = 9, N=3 for TRPC6-KD vs 1.90 ± 0.27 Hz, n = 9, N = 3 for scramble-shRNA, P < 0.01, Fig. 5E), and hyperpolarization of the RMP (-58.21 ± 1.79 mV, n = 10, N = 6 for TRPC6-KD vs -50.02± 0.67 mV, n = 10, N = 6 scramble-shRNA, P < 0.001, Fig. 5G) of the male VTA DA neurons. In addition, the inhibitory effect of 2-APB and LA on the firing of VTA DA neurons infected with TRPC6-shRNA largely diminished, with no statistical difference in firing frequency before and after dosing (2-APB: from 0.58 ± 0.15 Hz to 0.40 ± 0.15 Hz, n = 7, N = 5, P > 0.05, Fig. 5F; LA: from 0.43 ± 0.05 Hz to 0.43 ± 0.06 Hz, n = 8, N = 5, P > 0.05, Fig. 5H). While the above results were observed in the male mice, similar results were observed in the female mice (sFig.8 A and B).

Selective knockdown of TRPC6 in the mPFC-projecting VTA DA neurons produces depression-like and anxiety-like behaviors.

Earlier version: A. shRNA against TRPC6 carried by retro-AAV-loxp virus (TRPC6-shRNA) was injected into the mPFC of the DAT-Cre male mice. B. The mRNA level in the shRNA-TRPC6 transfected VTA (TRPC6-rcKD) DA neurons and the scramble shRNA transfected VTA DA neurons (Con) was analyzed using single-cell quantitative PCR (scqPCR) (Con: n = 8, N = 4; TRPC6-rcKD: n = 10, N = 4). Two-sample T test, t = 4.506, P = 0.0004. C. Immunofluorescence labelling showing the expression of AAV-retro-hSyn-DIO-mCherry-shRNA (TRPC6) (red) and DAT (green) in the VTA of the DAT-Cre mice (scale bar, 10 µm). D. The effects of Cre-induced conditional knockdown of TRPC6 in mPFC-projecting VTA DA neurons on the behaviors of mice: (i) the sucrose preference test, (ii) the tail suspension test, and (iii) the elevated plus-maze test (Con: N = 8; TRPC6-rcKD: N = 10). Two-sample T test, sucrose preference test: t = 5.215, P = 0.0001; tail suspension test: t = 2.993, P = 0.0097; elevated plus-maze test: open arm t = 5.802, P = 0.0001. ** P < 0.01, *** P < 0.001. n is number of neurons recorded and N is the number of mice used.

Revised version: A. shRNA against TRPC6 carried by retro-AAV-loxp virus (TRPC6-shRNA) was injected into the mPFC of the DAT-Cre male mice. B. The mRNA level in the shRNA-TRPC6 transfected mPFC-projecting VTA (TRPC6-rcKD) DA neurons and the scramble shRNA transfected mPFC-projecting VTA DA neurons (Con) was analyzed using single-cell quantitative PCR (scqPCR)(Con: n = 15, N = 6; TRPC6-rcKD: n = 16, N = 6). Mann-Whitney U test, U =8.500, ****P < 0.0001.C. Immunofluorescence labelling showing the expression of AAV-retro-hSyn-DIO-mCherry-shRNA(TRPC6) (red) and DAT (green) in the VTA of the DAT-Cre mice (scale bar, 10 µm). D. The effects of Cre-induced conditional knockdown of TRPC6 in mPFC-projecting VTA DA neurons on the behaviors of mice: (i) the sucrose preference test, (ii) the tail suspension test, and (iii) the elevated plus-maze test (Con: N = 10; TRPC6-rcKD: N = 10). Two-sample T test, sucrose preference test: t = 3.592, df = 18, 95% CI: −11.48 to −3.008, **P = 0.0021; tail suspension test: t = 3.707, df = 18, 95% CI: 10.79 to 39.01, **P = 0.0016; elevated plus-maze test: open arm t = 3.756, df = 18, 95% CI: −8.099 to −2.288, **P = 0.0014. **P < 0.01, ****P < 0.0001.n is number of neurons recorded and N is the number of mice used.

Removed in the revised version: Consequence of knocking down TRPC6 described above should not be a result of a secondary effect on other TRPC channels with which TRPC6 is known to interact to form heteromers, such as TRPC4 and TRPC7[47, 48], since in experiments using single cell PCR (sFig. 9A), it was found only a very small proportion of TRPC6-positive DA cells (DAT+) expressed TRPC4 (sFig. 9Bi) or TRPC7 (sFig. 9Bii), in consistent with the results of single-cell RNA-seq results (Fig. 2). Taken together, above results indicate TRPC6 is an important contributor in subthreshold depolarization and spontaneous firing of the VTA DA neurons.

Removed in the revised version: Down-regulation of TRPC6 contributes to the altered firing activity of the VTA DA neurons in depression model

Removed in the revised version: In multiple depression models, the depression-like behavior was directly linked to the altered firing activity of the VTA DA neurons[4, 7, 49–51]. In consideration of the evidence we described above that NALCN and TRPC6 play key roles in firing activity of the VTA DA neurons, we went further to study if these channels also contributed to the altered firing activity of the VTA DA neurons and to the development of the depression-like behavior in depression models of mice.

Removed in the revised version: We first established a mice depression model of chronic mild unpredictable stress (CMUS) [52], which, after 5 weeks subjecting to two different stressors every day (Fig. 6A), manifested chronic-stress-induced depression-like behaviors in the sucrose preference test (SPT) and the tail suspension test (TST), with reduced sucrose preference and lengthened immobility time, respectively (Fig. 6B). Multiple other behaviors tests were also performed on these CMUS mice (sFig. 10), all indicating depression/anxiety-like behaviors.

Earlier version: It has been shown that a reduced firing activity in the VTA DA neurons is responsible for the chronic-stress-induced depression-like behavior in the CMUS model[51, 53]. Consistent with this finding, we also found the firing frequency of the VTA DA neurons from the CMUS male mice was significantly reduced as compared with that from the control male mice (2.83 ± 0.42 Hz, n = 18, N = 5 for control mice vs 1.30 ± 0.14 Hz, n = 28, N = 5 for CMUS mice, P < 0.01, Fig. 6C). In agreement with a role for TRPC6 in this reduced firing activity, TRPC6 protein in the male VTA tissue was found to be down-regulated in western blot experiments (Fig. 6D). On the other hand, NALCN protein in the VTA was not altered in the CMUS male mice model of depression (sFig. 11).

Revised version: It has been shown that a reduced firing activity in the VTA DA neurons is responsible for the chronic-stress-induced depression-like behavior in the CMUS model[53, 55]. Consistent with this finding, we also found the firing frequency of the VTA DA neurons from the CMUS male and female mice was significantly reduced as compared with that from the control male and female mice (2.90 ± 0.17 Hz, n = 53, N = 31 for control mice vs 1.24 ± 0.08 Hz, n = 68, N = 33 for CMUS mice, Mann-Whitney U test, U = 415, P < 0.0001, Fig. 6C). In agreement with a role for TRPC6 in this reduced firing activity, TRPC6 protein in the VTA tissue of male and female CMUS mice was found to be down-regulated in western blot experiments (Fig. 6D). On the other hand, NALCN protein in the VTA was not altered in the CMUS male mice model of depression (sFig. 9).

Earlier version: We also tested if this down-regulation of TRPC6 protein could also be found in a similar but different model of depression. For this, a chronic restraint stress (CRS) model was used. After 3 weeks of restraint stress stimulation, the CRS male mice developed similar depression/anxiety-like behavior in the behavior tests like these similarly performed in the CMUS model (sFig. 12A-I). Importantly, male mice with CRS-induced depression also showed a significant downregulation of TRPC6 protein in the VTA (0.43 ± 0.09 fold, N = 6 for control and N = 9 for CRS, P < 0.001, sFig. 12J). In female CMUS mice (sFig13. A-I), TRPC6 protein in the VTA tissue was also found to be down-regulated in western blot experiments (sFig. 13J).

Revised version: We also tested if this down-regulation of TRPC6 protein could also be found in a similar but different model of depression. For this, a chronic restraint stress (CRS) model was used. After 3 weeks of restraint stress stimulation, the CRS male mice developed similar depression/anxiety-like behavior in the behavior tests like these similarly performed in the CMUS model (sFig. 10A-I). Importantly, male mice with CRS-induced depression also showed a significant downregulation of TRPC6 protein in the VTA (0.38 ± 0.06 fold, N = 10 for control and N = 14 for CRS, Two-sample T test, t = 7.369, df = 22, 95% CI: −0.8006 to −0.4489, P < 0.0001, sFig. 10J).

Earlier version: It is known that CMUS mainly decreases the activity of VTA dopamine neurons that project to mPFC[11, 51]. Assuming that the decreased firing activity of the mPFC-projecting VTA DA neurons in the CMUS mice was associated with the down-regulation of TRPC6 expression, we reasoned that the firing activity in these CMUS mice would respond less to the TRP channel inhibitors. Indeed, TRP channel inhibitor 2-APB and TRPC6 inhibitor LA decreased their inhibitory effect on the firing activity of the mPFC-projecting VTA DA neurons from the CMUS male mice (2-APB: 1.61 ± 0.23 Hz to 1.19 ± 0.22 Hz, n = 15, N=5, P > 0.05; LA: 0.97 ± 0.15 Hz to 0.66 ± 0.09 Hz, n = 19, N = 8, P > 0.05) compared with these from the control male mice (2-APB: 3.74 ± 0.42 Hz to 0.17 ± 0.09 Hz, n = 12, N = 5, P < 0.001; LA: 2.83 ± 0.17 Hz to 0.10 ± 0.07 Hz, n = 17, N = 8, P < 0.001) (Fig. 6E and F).

Revised version: It is known that the activity of VTA dopamine neurons that project to mPFC is decreased in CMUS mice with depressive behaviors[11, 53]. Assuming that the decreased firing activity of the mPFC-projecting VTA DA neurons in the CMUS mice was associated with the down-regulation of TRPC6 expression, we reasoned that the firing activity in these CMUS mice would respond less to the TRP channel inhibitors. Indeed, TRP channel inhibitor 2-APB and TRPC6 inhibitor LA decreased their inhibitory effect on the firing activity of the mPFC-projecting VTA DA neurons from the CMUS male and female mice (2-APB:1.56 ± 0.15Hz to 1.19 ± 0.14 Hz, n = 25, N = 15, Kruskal-Wallis test, P > 0.9999; LA: 1.12 ± 0.11Hz to 0.83 ± 0.08 Hz, n = 30, N = 18, Kruskal-Wallis test, P > 0.9999) compared with these from the control mice (2-APB: 3.35 ± 0.29 Hz to 0.17 ± 0.08 Hz, n = 22, N = 15, Kruskal-Wallis test, P < 0.0001; LA: 2.97 ± 0.16 Hz to 0.17 ± 0.08 Hz, n = 25, N = 16, Kruskal-Wallis test, P < 0.0001) (Fig. 6E and F).

Earlier version: Taken together, above results suggest down-regulation of TRPC6 is a key determinant for the altered firing activity of the VTA DA neuron in mice models of depression

Revised version: Taken together, above results suggest down-regulation of TRPC6 results in the altered firing activity of the VTA DA neuron in mice models of depression.

Down-regulation of TRPC6 in the VTA DA neurons confer the mice with depression-like behavior

Earlier version: If down-regulation of TRPC6 is crucial for the depression-like behaviors, like that indicated above in the CMUS and the CRS depression models, then down regulation of TRPC6 in non-stressfully treated mice should also develop depression-like behavior. We tested this possibility by selectively knocking down TRPC6 in the VTA DA neurons, using DAT-Cre male mice injected with a AAV viral construct (AAV9-hSyn-DIO-shRNA (TRPC6) -RFP) (Fig. 7A), driving the expression of shRNA against TRPC6 selectively in the VTA DA neurons. Three weeks after the AAV injection, the qPCR (Fig. 7B) and the immunofluorescence results (Fig. 7C) indicated an efficient down-regulation of TRPC6 in the VTA DA neurons.

Revised version: If down-regulation of TRPC6 is crucial for the depression-like behaviors, like that indicated above in the CMUS and the CRS depression models, the down regulation of TRPC6 in non-stressfully treated mice should also develop depression-like behavior. We tested this possibility by selectively knocking down TRPC6 in the VTA DA neurons, using DAT-Cre male mice injected with a AAV viral construct (AAV9-hSyn-DIO-shRNA (TRPC6) -RFP) (Fig. 7A), driving the expression of shRNA against TRPC6 selectively in the VTA DA neurons. Three weeks after the AAV injection, the qPCR (Fig. 7B) and the immunofluorescence results (Fig. 7C) indicated an efficient down-regulation of TRPC6 in the VTA DA neurons.

Earlier version: We next compared the depression-and anxiety-like behaviors in these conditionally TRPC6-knockout male mice (TRPC-cKD) with that in the control mice (infected with virus carrying the scrambled shRNA). As shown in Fig. 7D, in the TRPC6-cKD mice, the sucrose preference was significantly reduced (7Di, 72.60 ± 3.52%, N = 13 for TRPC6-cKD, vs 92.74 ± 0.86 %, N = 9 for control, P < 0.001), the immobility time in the TST was significantly lengthened (7Dii, 192.10 ± 3.13 s vs 133.20 ± 8.89 s, P < 0.001). In the elevated plus maze test, the percentage of time spent in the open arm was significantly reduced (7Diii, 2.43 ± 0.48 % vs 14.78 ± 1.26 %, P < 0.001) and the percentage of time spent in the closed arm was also significantly increased (7Div). And the Cre-dependent TRPC6-overexpression in male VTA DAT-Cre mice (TRPC6-cOE) rescued the selective-TRPC6-knockdown-induced depression-like behaviors (Fig. 7E-H). Furthermore, selectively knockdown of TRPC6 in the mPFC-projecting VTA DA neurons (virus-retro carrying DIO-TRPC6-shRNA (TRPC6-rcKD) was injected into the mPFC of DAT-Cre male mice) (Fig. 8A-C) also conferred the mice with the depression-and anxiety-like behaviors (Fig. 8D).

Revised version: We next compared the depression- and anxiety-like behaviors in these conditionally TRPC6-knockdown male mice (TRPC-cKD) with that in the control mice (infected with virus carrying the scrambled shRNA). As shown in Fig. 7D, in the TRPC6-cKD male mice, the sucrose preference was significantly reduced (7Di, 74.52 ± 2.00 %, N = 10 for TRPC6-cKD, vs 85.87 ± 1.85 %, N = 10 for control, Two-sample T test, t = 4.172, df = 18, 95% CI: −17.07 to −5.634, P = 0.0006), the immobility time in the Tail suspension test was significantly lengthened (7Dii, 142.00 ± 4.66 s vs 121.40 ± 3.94 s, Two-sample T test, t = 3.377, df = 18, 95% CI: 7.785 to 33.41, P = 0.0034). In the elevated plus maze test, the percentage of time spent in the open arm was significantly reduced (7Diii, 1.60 ± 0.50 % vs 11.42 ± 2.04 %, Mann-Whitney U test, U = 2, P < 0.0001) and the percentage of time spent in the closed arm was also significantly increased (7Div). And the Cre-dependent TRPC6-overexpression in male VTA DAT-Cre mice (TRPC6-cOE) rescued the selective-TRPC6-knockdown-induced depression-like behaviors (Fig. 7E-H). Furthermore, selectively knockdown of TRPC6 in the mPFC-projecting VTA DA neurons (virus-retro carrying DIO-TRPC6-shRNA (TRPC6-rcKD) was injected into the mPFC of DAT-Cre male mice) (Fig. 8A-C) also conferred the mice with the depression- and anxiety-like behaviors (Fig. 8D).

Earlier version: Taken together, these results suggest downregulation of TRPC6 in the VTA DA neurons confers the mice with phenotypes of depression/anxiety.

Revised version: Taken together, these results suggest down-regulation of TRPC6 in the VTA DA neurons promotes the phenotypes of depression/anxiety of the mice in the CMUS.

Discussion

In this study, we made a systematic study on the molecular mechanism for the subthreshold depolarization that drive spontaneous firing of the VTA DA neurons. We identified TRPC6 channels, alongside NALCN, as major contributors to this subthreshold depolarization and related spontaneous firing, and further importantly, we also demonstrated that TRPC6 contributed to the altered firing activity of the VTA DA neurons under states of chronic-stress-induced depression-like behaviors.

Earlier version: The nature of a relatively depolarized resting membrane potential in the VTA DA neurons imposes a need for a full understanding of the channel conductance underlying this depolarization. It is first clear from these current and previous studies that this conductance is mediated by a persistent Na+ influx. Unlike the DA neurons in the SNc, where Ca2+ influx is needed for the spontaneous firing [54], the results in our and others’ studies indicate Ca2+ influx are not necessary for the spontaneous firing of the VTA DA neurons, rather, replacement of Ca2+ ions in the extracellular fluid with Mg2+ accelerates the frequency of spontaneous firing (Fig. 1B) [15], possibly due to a reduced activity of the Ca2+-activated K+ currents and the increased activity of NALCN[28, 55]. Consistent with these finding and more precisely for a role of Ca2+ influx on the subthreshold depolarization, removing extracellular Ca2+ did not hyperpolarize the RMP. On the other hand, consistent the previous study [15], a TTX-insensitive Na+ influx clearly contributed to the subthreshold depolarization (Fig. 1D).

Revised version: The nature of a relatively depolarized resting membrane potential in the VTA DA neurons imposes a need for a full understanding of the channel conductance underlying this depolarization. It is first clear from these current and previous studies that this conductance is mediated by a persistent Na+ influx. Unlike the DA neurons in the SNc, where Ca2+ influx is needed for the spontaneous firing [56], the results in our and others’ studies indicate that Ca2+ influx are not necessary for the spontaneous firing of the VTA DA neurons. Replacement of Ca2+ ions in the extracellular fluid with Mg2+ accelerates the frequency of spontaneous firing (Fig. 1B) [15], possibly resulting from a reduced activity of the Ca2+-activated K+ currents and an increased activity of NALCN[28, 57]. Consistent with this, removing extracellular Ca2+ did not hyperpolarize the RMP, which verifies more precisely for a role of Ca2+ influx on the subthreshold depolarization. On the other hand, consistent the previous study [15], a TTX-insensitive Na+ influx clearly contributed to the subthreshold depolarization (Fig. 1D).

Earlier version: A TTX-insensitive nature of this Na+ conductance suggests that cation conductance permeable to Na+ might be the molecular candidates. We first focused on the HCN channels because, (1) it has been suggested HCN channels are important regulators for the spontaneous firing of VTA DA neurons under states of depression-like behaviors[7, 11]; (2) HCN are one of the more highly expressed NSCCs in the VTA DA neurons in our single cell RNA-Seq study; (3) HCN are permeable to Na+[56]. However, it is clear from our study that HCN does not contribute to the spontaneous firing of the VTA DA neurons in the adult mice; none of the known HCN blockers we used had any significant effects on the firing activity. Presence of HCN has long been used as a signature of DA neurons[57], also as a differentiating factor for the projection-specificity of the VTA DA neurons[39]. Nonetheless, contribution of HCN to the spontaneous firing of the VTA DA neurons has been controversial; the HCN blockers ZD7288 and CsCl have been reported as both effective[23, 24, 58] and non-effective [15] on the spontaneous firing of the VTA DA neurons. A closer inspection into our and others’ results present a possible explanation for these different results: in the younger rodents (less than one month old) (sFig.5) [59], the HCN in the DA neurons play a more important role whereas in the adult mice (∼ 8 weeks)[59], like the ones we used in this study, the HCN do not contribute to the spontaneous firing of the VTA DA neurons. This is likely due to a shift of activation gating of HCN channels in a hyperpolarization direction with age, thus the modulation of HCN on the firing activity diminishes with age, as it is described in the SNc DA neurons[59]. Consistent with this, we found hyperpolarizing the RMP rendered the VTA DA neurons of the adult mice with sensitivity to the HCN blockers. Thus, the VTA DA neurons in adult mice do not use HCN to depolarize the membrane to generate action potentials.

Revised version: A TTX-insensitive nature of this Na+ conductance suggests that cation conductance permeable to Na+ might be the molecular candidates. We first focused on the HCN channels because, (1) it has been suggested HCN channels are important regulators for the spontaneous firing of VTA DA neurons in adult male mice under states of depression-like behaviors[7, 11]; (2) HCN are one of the more highly expressed NSCCs in the VTA DA neurons in our single cell RNA-Seq study; (3) HCN are permeable to Na+[58]. However, it is clear from our study that HCN does not contribute to the spontaneous firing of the VTA DA neurons in the adult male mice; none of the known HCN blockers we used had any significant effects on the firing activity. Presence of HCN has long been used as a signature of DA neurons[59], also as a differentiating factor for the projection-specificity of the VTA DA neurons[40]. Nonetheless, contribution of HCN to the spontaneous firing of the VTA DA neurons has been controversial; the HCN blockers ZD7288 and CsCl have been reported as both effective[23, 24, 60] and non-effective [15] on the spontaneous firing of the VTA DA neurons. A closer inspection into our and others’ results present a possible explanation for these different results: in the younger rodents (less than one month old [61] and less than 15 days old in our study), the HCN in the DA neurons play a more important role whereas in the adult mice (∼ 8 weeks [61]), like the male adult ones (6∼8 weeks) we used in this study, the HCN do not contribute to the spontaneous firing of the VTA DA neurons. This is likely due to a shift of activation gating of HCN channels in a hyperpolarization direction with age, thus the modulation of HCN on the firing activity diminishes with age, as it is described in the SNc DA neurons[61]. Consistent with this, we found hyperpolarizing the RMP rendered the VTA DA neurons of the adult male mice with sensitivity to the HCN blockers. Thus, the VTA DA neurons in adult male mice do not use HCN to depolarize the membrane to generate action potentials.

Earlier version: NALCN is widely expressed in central neurons[60, 61], and is a popular candidate considered for the ion mechanism of the subthreshold depolarization currents and Na+ leaky currents. NALCN has been shown to be an important component of background Na+ currents and to be important for neuronal excitability in the hippocampal neurons, posterior rhomboid nucleus (Retrotrapezoid nucleus/RTN)) chemo-sensitive neurons (CO2/H+-sensitive neurons), GABA neurons and DA neurons in the substantia nigra, and spinal projection neurons [26, 28–31, 42]. Our results with single cell RNA-Seq and immunofluorescence methods validated high level expression of NALCN in the VTA DA neurons; both the pharmacological and the NALCN-gene-knockout experiments demonstrated convincingly involvement of NALCN in subthreshold depolarization and spontaneous firing of the VTA DA neurons.

Revised version: NALCN is widely expressed in central neurons[62, 63], and is a popular candidate considered for the ion mechanism of the subthreshold depolarization currents and Na+ leaky currents. NALCN has been shown to be an important component of background Na+ currents and to be important for neuronal excitability in the hippocampal neurons, posterior rhomboid nucleus (Retrotrapezoid nucleus/RTN) chemo-sensitive neurons (CO2/H+-sensitive neurons), GABA neurons and DA neurons in the substantia nigra, and spinal projection neurons [28–32, 44]. Our results with single cell RNA-Seq, single-cell PCR and immunofluorescence methods validated high level expression of NALCN in the VTA DA neurons; both the pharmacological and the NALCN-gene-knockout experiments demonstrated convincingly involvement of NALCN in subthreshold depolarization and spontaneous firing of the VTA DA neurons.

We did not detect alteration of NALCN protein expression in the depression model of CMUS in the VTA tissue. However, this does not completely exclude a possible contribution of NALCN to the altered functional activity of the VTA DA neurons which underlies the chronic-stress-induced depression-like behaviors. Future study focusing on the functional property of NALCN in the cellular level of the VTA DA neurons under a state of depression is needed to clarify this issue. NALCN is after all a potential drug target in consideration of its important contribution to the resting membrane potential and functional activity of neurons.

Earlier version: The most interesting and important finding of this study is the role of TRPC6 in regulation of the firing activity of the VTA DA neurons in both physiological and depression-state conditions, and its involvement in depression-like behaviors; the experimental evidence from pharmacological, gene silencing, electrophysiological and behavior results unambiguously support this conclusion. It is interesting to note, although with an unclear neuronal circuit mechanism of action, the TRPC6 channel opener hyperforin was previously reported to have antidepressant effects in corticosterone-depressed mice, which were reversed by prior administration of the TRPC6 blocker larixyl acetate[62]. Also, in a CMUS rat model, TRPC6 expression was found to be downregulated in the hippocampus, and administration of the TRPC6 opener hyperforin was effective in alleviating depression-like behavior[63]. Our study provides strong and convincing evidence that TRPC6 is a key regulator of the VTA DA neurons which are known to be a key player in depression states [7, 21, 49, 50, 64].

Revised version: The most interesting and important finding of this study is the role of TRPC6 in regulation of the firing activity of the VTA DA neurons in both physiological and depression-state conditions, and its involvement in depression-like behaviors; the experimental evidence from pharmacological, gene silencing, electrophysiological and behavior results unambiguously support this conclusion. It is interesting to note, although with an unclear neuronal circuit mechanism of action, the TRPC6 channel opener hyperforin was previously reported to have antidepressant effects in corticosterone-depressed mice, which were reversed by prior administration of the TRPC6 blocker larixyl acetate[64]. Also, in a CMUS rat model, TRPC6 expression was found to be downregulated in the hippocampus, and administration of the TRPC6 opener hyperforin was effective in alleviating depression-like behavior[65]. Our study provides strong and convincing evidence that TRPC6 is a key regulator of the VTA DA neurons which are known to be a key player in depression states [7, 21, 51, 52, 66].

Earlier version: The above-mentioned role for TRPC6 should come from its contribution to the subthreshold depolarization, through a persistent permeation to the influx of Na+. Although TRPC channels has been better recognized for their permeability to Ca2+ and their regulation on cellular Ca2+ homeostasis, the importance of TRPC channels for the permeability of monovalent cations, especially Na+, is now gradually being appreciated [65]. In the substantial nigra GABA and DA neurons, TRPC3 has been reported to play an important role in maintaining depolarization membrane potential, pacemaking, and firing regularity[31, 33]. Otherwise, it has been reported that slow tonic firing of SNc DA neurons, depends on the basal activity of both the NALCN and TRPC3 channels, but that burst firing does not require TRPC3 channels but relies only on NALCN channels[66]. TRPC6 and TRPC3 belong to the same TRPC family subclass and TRPC6 is more than 75% homologous to TRPC3[45]. Interestingly, our results with single-cell RNA-Seq and single-cell PCR demonstrated the rich expression of TRPC6 in the VTA DA neurons among all TRP channels but no expression of TRPC3 was found in these neurons. These results present an opportunity for selectively targeting a single TRPC channel to exert selective pharmacological results. Besides, TRPC4 expression was detected in a uniformly distributed subset of VTA DA neurons[35], which is consistent with our results with single-cell RNA-Seq, and TRPC4 KO rats showed decreased VTA DA neuron tonic firing and deficits in cocaine reward and social behaviors[35].

Revised version: The above-mentioned role for TRPC6 should come from its contribution to the subthreshold depolarization, through a persistent permeation to the influx of Na+. Although TRPC channels has been better recognized for their permeability to Ca2+ and their regulation on cellular Ca2+ homeostasis, the importance of TRPC channels for the permeability of monovalent cations, especially Na+, is now gradually being appreciated [67]. In the substantial nigra GABA and DA neurons, TRPC3 has been reported to play an important role in maintaining depolarization membrane potential, peacemaking, and firing regularity[32, 34]. Otherwise, it has been reported that slow tonic firing of SNc DA neurons, depends on the basal activity of both the NALCN and TRPC3 channels, but that burst firing does not require TRPC3 channels but relies only on NALCN channels[68]. TRPC6 and TRPC3 belong to the same TRPC family subclass and TRPC6 is more than 75% homologous to TRPC3[47]. Interestingly, our results with single-cell RNA-Seq and single-cell PCR demonstrated the rich expression of TRPC6 in the VTA DA neurons among all TRP channels but no expression of TRPC3 was found in these neurons. These results present an opportunity for selectively targeting a single TRPC channel to exert selective pharmacological results. Besides, TRPC4 expression was detected in a uniformly distributed subset of rat VTA DA neurons[36], which differs with our results with single-cell RNA-Seq and single-cell PCR, and TRPC4 KO rats showed decreased VTA DA neuron tonic firing and deficits in cocaine reward and social behaviors[36].TRPC4 is detected to be expressed with a low level in the VTA DA neurons of male mice by single-cell RNA-Seq (Fig.2) and in TRPC6 and DAT double-positive VTA cells by single-cell PCR (sFig.6).The different species used in this and our study may be one explanation, but the role of TRPC4 in the excitability of mice VTA TRPC6-negative DA neurons is worth to investigating further.

Earlier version: It has also been demonstrated in cell types other than neurons, such as rat portal vein smooth muscle cells, the activation of α-adrenergic receptor can open TRPC6, which leads to cell depolarization, activation of voltage-dependent Ca2+ channels, and ultimately the contraction of smooth muscle cell[67]; it is also shown that angiotensin II and endothelin I can activate TRPC3/C6 heterodimers via vascular G protein-coupled receptors (GPCR), thereby depolarizing cells [68, 69]. Following this line, it would be very interesting to know if TRPC6 in the VTA DA neurons, where multiple GPCR reside receiving input and auto modulation by neuronal transmitters [70, 71], could also be a target of modulation with similar mechanism. In consideration of the important role, we described here in this study, this type of modulation will present possible explanations for some unsolved mechanism for physiological and pathophysiological functions of the VTA DA neurons.

Revised version: It has also been demonstrated in cell types other than neurons, such as rat portal vein smooth muscle cells, the activation of α-adrenergic receptor can open TRPC6, which leads to cell depolarization, activation of voltage-dependent Ca2+ channels, and ultimately the contraction of smooth muscle cell[69]; it is also shown that angiotensin II and endothelin I can activate TRPC3/C6 heterodimers via vascular G protein-coupled receptors (GPCR), thereby depolarizing cells [70, 71]. Following this line, it would be very interesting to know if TRPC6 in the VTA DA neurons, where multiple GPCR reside receiving input and auto modulation by neuronal transmitters [72, 73], could also be a target of modulation with similar mechanism. In consideration of the important role, we described here in this study, this type of modulation will present possible explanations for some unsolved mechanism for physiological and pathophysiological functions of the VTA DA neurons.

Earlier version: The facts that down-regulation of TRPC6 proteins was correlated with reduced firing activity of the VTA DA neurons, the chronic-stress-induced depression-like behaviors, and that knocking down of TRPC6 in the VTA DA neurons confers the mice with depression-like behaviors strongly suggest a crucial role for TRPC6 in the development of chronic-stress-induced depression-like behaviors under stressful conditions like CMUS. To reinforce this, down-regulation of TRPC6 was also found in another chronic-stress-induced depression model of CRS, a model with similar neuronal alteration of reduced firing activity in the VTA DA neurons to that described in the CMUS model[72]. These facts with the fact that TRPC6 activator hyperforin is an effective antidepressant in multiple depression models[62, 63], and that hyperforin is the principal component of St. John’s wort, a well-known antidepressant herb[73], present TRPC6 as a very attractive drug target for new lines of antidepressants.

Revised version: The facts that down-regulation of TRPC6 proteins was correlated with reduced firing activity of the VTA DA neurons as well as the chronic-stress-induced depression-like behaviors of mice with both sexes, and that knocking down of TRPC6 in the VTA DA neurons confers the male mice with depression-like behaviors strongly suggest a crucial role for TRPC6 in the development of chronic-stress-induced depression-like behaviors under stressful conditions like CMUS. To reinforce this, down-regulation of TRPC6 was also found in another chronic-stress-induced depression model of CRS, a model with similar neuronal alteration of reduced firing activity in the VTA DA neurons to that described in the CMUS model[74]. These facts with the fact that TRPC6 activator hyperforin is an effective antidepressant in multiple depression models[64, 65], and that hyperforin is the principal component of St. John’s wort, a well-known antidepressant herb[75], present TRPC6 as a very attractive drug target for new lines of antidepressants.

Materials and methods

Animal Preparation

Earlier version: Male or female 6-8-week-old C57BL/6 (Vital River, China) and DAT-Cre mice with a C57BL/6 background (Stock No: 006660, the Jackson Laboratories, USA) were used for the studies. All experiments were conducted in accordance with the guidelines of Animal Care and Use Committee of Hebei Medical University and approved by the Animal Ethics Committee of Hebei Medical University.

Revised version: Male or female 6-8-week-old C57BL/6 (Vital River, China) and DAT-Cre mice with a C57BL/6 background (Stock No: 006660, the Jackson Laboratories, USA) were used for the studies. Ethics Statement: All experiments were conducted in accordance with the guidelines of Animal Care and Use Committee of Hebei Medical University and approved by the Animal Ethics Committee of Hebei Medical University (Shijiazhuang, China), 2017013.

Brain slice preparation

Earlier version: The details for the preparation of coronal brain slice containing VTA were the same as described in our previously published work[21, 70]. Briefly, mice were anesthetized with chloral hydrate (200 mg/kg, i.p.). After intracardial perfusion with ice-cold sucrose solution, the brains of the mice were removed quickly and placed into the slicing solution. The ice-cold sucrose-cutting solution contained (in mM): sucrose 260, NaHCO3 25, KCl 2.5, NaH2PO4 1.25, CaCl2 2, MgCl2 2, and D-glucose 10; osmolarity, 295-305 mOsm). Using the vibratome (VT1200S; Leica, Germany), the coronal midbrain slices (250 µm-thick) containing VTA were sectioned. The slices were incubated for 30 min at 36°C in oxygenated artificial cerebrospinal fluid (ACSF) (in mM: NaCl 130, MgCl2 2, KCl 3, NaH2PO4-2H2O 1.25, CaCl2 2, D-Glucose 10, NaHCO3 26; osmolarity, 280-300 mOsm), and were then left for recovery for 60 min at room temperature (23-25°C) until use.

Revised version: The details for the preparation of coronal brain slice containing VTA were the same as described in our previously published work[21, 72]. Briefly, mice were anesthetized with chloral hydrate (200 mg/kg, i.p.). After intracardial perfusion with ice-cold sucrose solution, the brains of the mice were removed quickly and placed into the slicing solution. The ice-cold sucrose-cutting solution contained (in mM): sucrose 260, NaHCO3 25, KCl 2.5, NaH2PO4 1.25, CaCl2 2, MgCl2 2, and D-glucose 10; osmolarity, 295-305 mOsm). Using the vibratome (VT1200S; Leica, Germany), the coronal midbrain slices (250 µm-thick) containing VTA were sectioned. The slices were incubated for 30 min at 36°C in oxygenated artificial cerebrospinal fluid (ACSF) (in mM: NaCl 130, MgCl2 2, KCl 3, NaH2PO4 1.25, CaCl2 2, D-Glucose 10, NaHCO3 26; osmolarity, 280-300 mOsm), and were then left for recovery for 60 min at room temperature (23-25°C) until use.

The brain slices were transferred to the recording chamber and were continuously perfused with fully oxygenated ASCF during recording.

Identification of DA neurons and electrophysiological recordings

Earlier version: Recordings in the slices were performed in whole-cell current-clamp and voltage-clamp configurations on the Axoclamp 700B preamplifier (Molecular Devices, USA) coupled with a Digidata 1550B AD converter (Molecular Devices, USA). Neurons in the VTA were visualized with a 40X water-immersion objective equipped by an optiMOS microscope camera (Qimaging, Canada) on an Olympus-BX51 microscope (Olympus, Japan). Projection-specific or GFP-positiveVTA neurons were identified by infrared-differential interference contrast (IR-DIC) video microscopy and epifluorescence (Olympus, Japan) for detection of retrobeads (red) positive or GFP-positive neurons.

Revised version: Recordings in the slices were performed in whole-cell current-clamp and voltage-clamp configurations on the Axoclamp 700B preamplifier (Molecular Devices, USA) coupled with a Digidata 1550B AD converter (Molecular Devices, USA). Neurons in the VTA were visualized with a 40× water-immersion objective equipped by an optiMOS microscope camera (Qimaging, Canada) on an Olympus-BX51 microscope (Olympus, Japan). Projection-specific or GFP-positive VTA neurons were identified by infrared-differential interference contrast (IR-DIC) video microscopy and epifluorescence (Olympus, Japan) for detection of retrobeads (red) positive or GFP-positive neurons.

Earlier version: For whole-cell recording, glass electrodes (3-5 MΩ) were filled with internal solution (in mM): K-methylsulfate 115, KCl 20, MgCl2 1, HEPES 10, EGTA 0.1, MgATP 2 and Na2GTP 0.3, pH adjusted to 7.4 with KOH. And the extracellular solution was the ACSF.

Revised version: For whole-cell recording, glass electrodes (3-5 MΩ) were filled with internal solution (in mM): K-methyl sulfate 115, KCl 20, MgCl2 1, HEPES 10, EGTA 0.1, MgATP 2 and Na2GTP 0.3, pH adjusted to 7.4 with KOH. And the extracellular solution was the ACSF.

In whole-cell current-clamp mode, HCN function was judged by the inwardly rectifying characteristic sag potential generated by giving a hyperpolarizing current (−100 pA).

When recording the effect of removing extracellular Ca2+ on spontaneous cell discharge, CaCl2 in the ACSF was replaced with MgCl2, and the final MgCl2 concentration was 4 mM.

The resting membrane potential (RMP) was measured in current clamp mode (I = 0); the composition of the recording solution (mM) contained: NaCl 151, KCl 3.5, CaCl2 2, MgCl2 1, glucose 10, and HEPES 10, and the pH was adjusted to about 7.35 with NaOH. 1 μM tetrodotoxin was added to the extracellular solution to abolish action potential during measurement of the RMP.

To observe the effect of extracellular Na+ on the RMP, NaCl in the original recording solution was replaced with equimolar NMDG, and the NMDG-recording solution (mM) contained: NMDG 151, KCl 3.5, CaCl2 2, MgCl2 1, glucose 10, and HEPES 10, and the pH was adjusted to 7.35 with KOH.

Earlier version: For recording spontaneous firing of the neurons, cell-attached “loose-patch” (100-300 MΩ) recordings were used[74]. In this case, patch-pipettes (2-4 MΩ) were filled with ACSF, and the spontaneous activity was recorded in the current-clamp mode (I = 0). The synaptic blockers (CNQX, 10 μM; APV, 50 μM and gabazine, 10 μM) were added to isolate the intrinsic firing properties.

Revised version: For recording spontaneous firing of the neurons, cell-attached “loose-patch” (100-300 MΩ) recordings were used[76]. In this case, patch-pipettes (2-4 MΩ) were filled with ACSF, and the spontaneous activity was recorded in the current-clamp mode (I = 0). The synaptic blockers (CNQX, 10 μM; APV, 50 μM and gabazine, 10 μM) were added to isolate the intrinsic firing properties.

At the end of electrophysiological recordings, the recorded VTA neurons were collected for single-cell PCR.VTA DA neurons were identified by single-cell PCR for the presence of TH and DAT.

Retrograde labelling

Earlier version: Briefly, under general chloral hydrate anaesthesia (200 mg/kg, i.p.) and stereotactic control (RWD Instruments, Guangzhou, China), the skull surface was exposed. All coordinates are relative to bregma in mm using landmarks and neuroanatomical nomenclature that was described in the Franklin and Paxinos mouse brain atlas[75]. Red retrobeads (100 nl for single injection; Lumafluor Inc., Naples, FL, USA) were injected into the following sites using KD scientific syringe pump (KD scientific, Holliston MA, USA): bilaterally into NAc core (NAc c), NAc lateral shell (NAc ls), NAc medial shell (NAc ms) and basolateral amygdala (BLA), 4 separate sites (2 per hemisphere) into medial prefrontal cortex (mPFC). Coordinates for infusions were as follows: NAc c (AP +1.50, LM 0.84, DV -4.0; 100 nl beads); NAc ls (AP +0.86, LM 1.72, DV -4.0; 100 nl beads); NAc ms (AP +1.70, LM 0.53, DV -4.0; 100 nl beads); mPFC (AP +2.05 and 2.15, LM 0.27, DV -2.1 + 1.7; 200 nl beads); BLA (AP -1.46, LM 2.85, DV -4.3; 100 nl beads). Retrobeads were delivered through a pulled glass pipette using a PAP107 Multipipette Puller (MicroData Instrument, Inc., USA) and at a rate of 100 nl/min; the injection needle was left in place for at least 5 minutes after each infusion. Following surgery, mice were returned to single housing. For sufficient labelling, survival periods for retrograde tracer transport depended on respective injection areas: NAc c, NAc ls, NAc ms, 14 days; mPFC, 21 days; BLA 14 days[39]. Coronal sections of injection sites were stained with 4, 6-diamidino-2-phenylindole (DAPI, Sigma, USA) to confirm representative target location. Then, serial analyses of the injection-sites were carried out routinely.

Revised version: Briefly, under general chloral hydrate anaesthesia (200 mg/kg, i.p.) and stereotactic control (RWD Instruments, Guangzhou, China), the skull surface was exposed. All coordinates are relative to bregma in mm using landmarks and neuroanatomical nomenclature that was described in the Franklin and Paxinos mouse brain atlas[77]. Red retrobeads (100 nl for single injection; Lumafluor Inc., Naples, FL, USA) were injected into the following sites by KD scientific syringe pump (KD scientific, Holliston MA, USA): bilaterally into NAc core (NAc c), NAc lateral shell (NAc ls), NAc medial shell (NAc ms) and basolateral amygdala (BLA), 4 separate sites (2 per hemisphere) into medial prefrontal cortex (mPFC). Coordinates for infusions were as follows: NAc c (AP +1.50, ML ±0.84, DV −4.00∼-3.80; 100 nl beads); NAc ls (AP +0.86, ML ±1.72, DV −4.24∼-3.93; 100 nl beads); NAc ms (AP +1.70, ML ±0.53, DV −4.16∼--3.60; 100 nl beads); mPFC (AP +2.05 and 2.15, ML ±0.27, DV −2.26∼-1.70; 200 nl beads); BLA (AP −1.46, ML ±2.85, DV −4.01∼-3.58; 100 nl beads). Retrobeads were delivered through a pulled glass pipette using a PAP107 Multipipette Puller (MicroData Instrument, Inc., USA) and at a rate of 100 nl/min; the injection needle was left in place for at least 5 minutes after each infusion. Following surgery, mice were returned to single housing. For sufficient labelling, survival periods for retrograde tracer transport depended on respective injection areas: NAc c, NAc ls, NAc ms, 14 days; mPFC, 21 days; BLA 14 days[40]. Coronal sections of injection sites were stained with 4, 6-diamidino-2-phenylindole (DAPI, Sigma, USA) to confirm representative target location. Then, serial analyses of the injection-sites were carried out routinely.

AAV for gene knockdown or overexpression and viral construct and injection

Earlier version: For knockdown of NALCN and TRPC6, AAV9-U6-shRNA (NALCN)-CMV-GFP (300 nl) or AAV9-U6-shRNA (TRPC6)-CMV-GFP or its control AAV9-scramble-shRNA was delivered into the VTA (AP, −3.08 mm; ML, ±0.50 mm; DV, −4. 50 mm) of the mice. AAV9-hSyn-DIO-shRNA (TRPC6)-RFP (300 nl) or its control AAV9-scramble-shRNA was delivered into the VTA of the DAT-Cre C57BL/6 mice. AAV-retro-hSyn-DIO-mCherry-shRNA (TRPC6) (300 nl) or its control AAV-retro-scramble-shRNA was delivered into the mPFC of the DAT-Cre C57BL/6 mice.

Revised version: For knockdown of NALCN and TRPC6, AAV9-U6-shRNA(NALCN)-CMV-GFP (300 nl) or AAV9-U6-shRNA(TRPC6)-CMV-GFP or its control AAV9-scramble-shRNA was delivered into the VTA (AP −3.08 mm; ML ±0.50 mm; DV −5.00∼-4.20 mm) of the mice. AAV9-hSyn-DIO-shRNA(TRPC6)-RFP (300 nl) or its control AAV9-scramble-shRNA was delivered into the VTA of the DAT-Cre C57BL/6 mice. AAV-retro-hSyn-DIO-mCherry-shRNA(TRPC6) (300 nl) or its control AAV-retro-scramble-shRNA was delivered into the mPFC of the DAT-Cre C57BL/6 mice.

Earlier version: The shRNA hairpin sequences used in this study: NALCN shRNA: AAGATCGCACAGCCTCTTCAT[26]; TRPC6 shRNA: 5’-CCAGGATCAATGCATACAA-3’.

Revised version: The shRNA hairpin sequences used in this study:

Added in the revised version: NALCN shRNA: AAGATCGCACAGCCTCTTCAT[29];

Added in the revised version: TRPC6 shRNA: 5’-CCAGGATCAATGCATACAA-3’.

Earlier version: For Cre-dependent overexpression of TRPC6, AAV9-DIO-TRPC6-GFP (300 nl) or its control AAV9-scramble-RNA was delivered into the VTA of the DAT-Cre C57BL/6 mice, 7 days after injection of AAV9-hSyn-DIO-shRNA (TRPC6)-RFP.

Revised version: For Cre-dependent overexpression of TRPC6, AAV9-DIO-TRPC6-GFP (300 nl) or its control AAV9-scramble-shRNA was delivered into the VTA of the DAT-Cre C57BL/6 mice, 7 days after injection of AAV9-hSyn-DIO-shRNA(TRPC6)-RFP.

Mice were singly housed with enough food and water to recover for 4-5 weeks after injection of virus, before behavior tests and electrophysiological recordings.

Single cell PCR

Earlier version: mRNA was reversely transcribed to cDNA by PrimeScriptTMⅡ1st Strand cDNA Synthesis Kit (Takara-Clontech, Kyoto, Japan). At the end of electrophysiological recordings, the recorded neuron was aspirated into a pipette and then expelled into a PCR sterile tube containing 1 µl oligo-dT Primer and 1 µl dNTP mixture. The mixture was heated to 65°C for 5 min to denature the nucleic acids and then cooled on ice for 2 min. Reverse transcription from mRNA into cDNA was performed at 50°C for 50 min and then 85°C for 5 sec. cDNA was stored at −40°C. Then two rounds of conventional PCR with pairs of gene-specific outside (first round) and inner primers (second round) for GAPDH (positive control), TH, DAT, D2, GIRK2, Vgult2, GAD1, NALCN, TRPC6, TRPV2 and TRPC3 using GoTaq Green Master Mix (Promega, Madison, USA) were performed. After adding the specific outside primer pairs into each PCR tube, first-round synthesis conditions were as follows: 95 °C (5 min); 30 cycles of 95 °C (50 s), 58-63 °C (50 s), 72 °C (50 s); 72 °C(5 min). Then, the product of the first PCR was added in the second amplification round by using specific inner primer (final volume 25 μl). The second amplification round consisted of the following: 95 °C (5 min); 35 cycles of 95 °C (50 s), 58 °C-62 °C (45 s), 72 °C (50 s) and 5 min elongation at 72 °C. The final PCR products were separated by electrophoresis on 2% agarose gels. The negative control reactions with no added template were also performed in each experiment.

Revised version: mRNA was reversely transcribed to cDNA by PrimeScriptTM Ⅱ 1st Strand cDNA Synthesis Kit (Takara-Clontech, Kyoto, Japan). At the end of electrophysiological recordings, the recorded neuron was aspirated into a pipette and then expelled into a PCR sterile tube containing 1 µl oligo-dT Primer and 1 µl dNTP mixture. The mixture was heated to 65°C for 5 min to denature the nucleic acids and then cooled on ice for 2 min. Reverse transcription from mRNA into cDNA was performed at 50°C for 50 min and then 85°C for 5 min. cDNA was stored at −40°C. Then two rounds of conventional PCR with pairs of gene-specific outside (first round) and inner primers (second round) for GAPDH (positive control), TH, DAT, D2, GIRK2, VGulT2, GAD1, NALCN, TRPC6, TRPV2, TRPC3, TRPC4 and TRPC7 using GoTaq® Green Master Mix (Promega, Madison, USA) were performed. After adding the specific outside primer pairs into one PCR tube per cell, first-round synthesis conditions were as follows: 95 °C (5 min); 30 cycles of 95 °C (50 s), 58-63 °C (50 s), 72 °C (50 s); 72 °C (5 min). Then, the product of the first PCR was added in the second amplification round by using specific inner primer (final volume 25 μl). The second amplification round consisted of the following: 95 °C (5 min); 35 cycles of 95 °C (50 s), 58 °C-62 °C (45 s), 72 °C (50 s) and 5 min elongation at 72 °C. The final PCR products were separated by electrophoresis on 2% agarose gels. The negative control reactions with no added cDNA were also performed in each experiment.

The “outer” primers (from 5′ to 3′) as follows:

GAPDH:

AAATGGTGAAGGTCGGTGTGAACG (sense)

AGTGATGGCATGGACTGTGGTCAT (antisense)

Earlier version: TH:

Revised version: TH: GCCGTCTCAGAGCAGGATAC

Removed in the revised version: GCCGTCTCAGAGCAGGATAC

GGGTAGCATAGAGGCCCTTC

DAT:

CTGCCCTGTCCTGAAAGGTGT

GCCCAGTGATCACAGACTCC

D2-R:

AGCATCGACAGGTACACAGC

CCATTCTCCGCCTGTTCACT

GIRK2:

TGGACCAGGATGTGGAAAGC

AAACCCGTTGAGGTTGGTGA

Added in the revised version: VGluT2:

Removed in the revised version: Vgult2:

CTGCTTCTGGTTGTTGGCTACTCT

ATCTCGGTCCTTATAGGTGTACGC

GAD1:

ACAACCTTTGGCTGCATGTGGATG

AATCCCACGGTGCCCTTTGCTTTC

NALCN:

TCCATCTGTGGGAAGCATGT

CAAAAGCTGGTCCTCTTCAGTG

TRPC6:

AGCCTGTCTATTGAGGAAGAAC

AGCGAGAATGATTGGGGTCA

TRPV2:

CTGCACATCGCCATAGAGAA

AGGCTGGTGGTAGGCAACTA

TRPC3:

GCTGGCCAACATAGAGAAGG

CCTGCACGTGACTATCCACA

TRPC4:

GTCTATGTAGGCGATGCGCT

AGGAGGTCCTTGGCAAATTGT

TRPC7:

GCCATCAGCAAGGGCTATGT

CACGCCCACCACAAAATCC

The “inner” primers (from 5′ to 3′) as follows:

GAPDH:

GCAAATTCAACGGCACAGTCAAGG

TCTCGTGGTTCACACCCATCACAA

TH:

AGGAGAGGGATGGAAATGCT

ACCAGGGAACCTTGTCCTCT

DAT:

ATTTTGAGCGTGGTGTGCTG

TGCCTCACAGAGACGGTAGA

D2-R:

CCATTGTCTGGGTCCTGTCC

GTGGGTACAGTTGCCCTTGA

GIRK2:

AGCCGAGACAGGACCAAAAG

ATGTACGCAATCAGCCACCA

Added in the revised version: VGluT2:

Removed in the revised version: Vgult2:

CATCTCCTTCTTGGTGCTTGCAGT

ACAGCGTGCCAACGCCATTTGAAA

GAD1:

AGTCACCTGGAACCCTC

GCTTGTCTGGCTGGAA

NALCN:

GCCTTTGCTGGAGTTGTTCTG

CCCTGCATAATTGCCACAGTC

TRPC6:

TGCTAGAAGAGTGTCATTCCCT

Earlier version: CCTCCACAATCCGTACATAACC TRPV2:

Revised version: CCTCCACAATCCGTACATAACC

Added in the revised version: TRPV2:

GTGGGATGTGGTGACCTACC

GCTGGTACAGCCCTGAGAAC

TRPC3:

CCTTGGGTCTTCCATTCCTC

CACAACTGCACGATGTACTCC

TRPC4:

CCACGAGGTCCGCTGTAAC

CTCCCAACTTAACTGAAAGGCA

TRPC7:

CGCTTCTCCCACGACATCAC

ACTGGATAGGGACAGGTAGGC

RT-qPCR

Earlier version: Total RNA was prepared using Trizol. RNA was reversely transcribed into cDNA by TAKARA PrimeScripttm RT reagent Kit with gDNA Eraser for Reverse transcription. Subsequently, the SYBR Green Master Mix (TaKaRa) was used to do the qRT-PCR assays.

Revised version: Total RNA was prepared using Trizol. RNA was reversely transcribed into cDNA by TAKARA PrimeScriptTM RT reagent Kit with gDNA Eraser for Reverse transcription. Subsequently, the TB Green® Premix Ex Taq™ II (Tli RNaseH Plus) (TaKaRa) was used to do the RT-qPCR assays.

The specific primers were:

GAPDH:

GCAAATTCAACGGCACAGTCAAGG

TCTCGTGGTTCACACCCATCACAA

TRPC6:

GCAGAAACACAGAGGAAGTGG

GCTCTTTCCAGCTTGGCATATC

NALCN:

TAATGAGATAGGCACGAGTA

TGATGAAGTAGAAGTAGGAG

Single cell RT-qPCR

Earlier version: Following the retrograde labelling and VTA brain slice preparation, projection-specific VTA neurons were identified by infrared-differential interference contrast (IR-DIC) video microscopy and epifluorescence (Olympus, Japan) for detection of retrobeads (red) positive neurons. mRNA was reversely transcribed to cDNA by PrimeScriptTMⅡ1st Strand cDNA Synthesis Kit (Takara-Clontech, Kyoto, Japan). The projection-specific neurons was aspirated into a pipette and then expelled into a PCR sterile tube containing 1 µl oligo-dT Primer and 1 µl dNTP mixture. The mixture was heated to 65°C for 5 min to denature the nucleic acids and then cooled on ice for 2 min. Reverse transcription from mRNA into cDNA was performed at 50°C for 50 min and then 85°C for 5 sec. cDNA was stored at −40°C. The targeted pre-amplification was done, with pairs of targeted primers for GAPDH, TH and TRPC6, to quantify multiple targets per cell. After adding the targeted primer pairs into each PCR tube, the synthesis conditions were as follows: 95 °C (5 min); 10 cycles of 95 °C (50 s), 58-62 °C (50 s), 72 °C (50 s); 72 °C (5 min). After targeted pre-amplification, Quantitative PCR (qPCR) is the final laboratory step of the scRT-qPCR workflow with the SYBR Green Master Mix (TaKaRa).

Revised version: Following the retrograde labelling and VTA brain slice preparation, projection-specific VTA neurons were identified by infrared-differential interference contrast (IR-DIC) video microscopy and epifluorescence (Olympus, Japan) for detection of retrobeads (red) positive neurons. mRNA was reversely transcribed to cDNA by PrimeScriptTM Ⅱ 1st Strand cDNA Synthesis Kit (Takara-Clontech, Kyoto, Japan). The projection-specific neuron was aspirated into a pipette and then expelled into a PCR sterile tube containing 1 µl oligo-dT Primer and 1 µl dNTP mixture. The mixture was heated to 65°C for 5 min to denature the nucleic acids and then cooled on ice for 2 min. Reverse transcription from mRNA into cDNA was performed at 50°C for 50 min and then 85°C for 5 min. cDNA was stored at −40°C. The targeted pre-amplification was done, with pairs of targeted primers for GAPDH, TH and TRPC6, to quantify multiple targets per cell. After adding the targeted primer pairs into each PCR tube, the synthesis conditions were as follows: 95 °C (5 min); 10 cycles of 95 °C (50 s), 58-62 °C (50 s), 72 °C (50 s); 72 °C (5 min). After targeted pre-amplification, quantitative PCR (qPCR) is the final laboratory step of the scRT-qPCR workflow with the TB Green® Premix Ex Taq™ II (Tli RNaseH Plus) (TaKaRa).

Earlier version: The targeted primers for pre-amplification (from 5′to 3′) as follows: GAPDH:

Revised version: The targeted primers for pre-amplification (from 5′to 3′) as follows:

Added in the revised version: GAPDH:

AAATGGTGAAGGTCGGTGTGAACG (sense)

AGTGATGGCATGGACTGTGGTCAT (antisense)

TH:

AAGGTTCATTGGACGGCGG

ACATCGTCAGACACCCGAC

TRPC6:

AGCCTGTCTATTGAGGAAGAAC

AGCGAGAATGATTGGGGTCA

The primers for qPCR (from 5′ to 3′) as follows:

GAPDH:

GCAAATTCAACGGCACAGTCAAGG

TCTCGTGGTTCACACCCATCACAA

TH:

TCTCCTTGAGGGGTACAAAACC

ACCTCGAAGCGCACAAAGT

TRPC6:

TGCTAGAAGAGTGTCATTCCCT

CCTCCACAATCCGTACATAACC

Immunofluorescence

After intracardial perfusion with 4% paraformaldehyde (PFA) in 0.01 M PBS (pH 7.4), the brains were post-fixed in 4% paraformaldehyde. 48 h later, the brain tissue was placed in 30% sucrose solution (PBS preparation) to dehydrate. The brain tissue was sectioned coronally, including the nucleus accumbens (NAc), medial prefrontal cortex (mPFC), basolateral amygdala (BLA), and VTA, using a vibrating microtome. The section thickness of NAc, mPFC and BLA was 80μm and the section thickness of VTA was 40 μm. The sections were placed in PBS solution and stored in a refrigerator at 4°C for storage.

The brain section was incubated in 0.3% triton/3% BSA for 1 h at room temperature and then was blocked with 10% donkey serum at 37°C for 1 h. After that, the brain section was incubated in the corresponding antibodies in PBS at 4℃ for 24 h. The section was washed three times (10 min) with PBS. Finally, the brain section was incubated in the secondary antibodies for 2 h at 37°C. Images were obtained on a Leica TCS SP5 confocal laser microscope (Leica, Germany) equipped with laser lines for 405 mm, 488 mm, 561 mm and 647 mm illumination. Images were analyzed with LAS-AF-Lite software (Leica, Germany).

Single-cell whole-transcriptome gene sequencing

Earlier version: After retrobeads injection, following brain slice preparation and 1 h recording, the recorded and retrobeads-labelled neurons were aspirated into the patch pipette and were then broken into the PCR tube containing 1 μl lysis buffer. For mRNA in individual cells, mRNA was amplified and cDNA by SMARTer Ultra Low Input RNA for Illumina Kit, which was qualified and reversely transcribed to cDNA by Qubit and Agilent Bioanalyzer 2100 electrophoresis. After fragmentation of cDNA (300 bp) by ultrasound, sequencing libraries (end repair, addition of poly(A), and ligation of sequencing connectors) were built using the Ovation Ultralow Library System V2. After that, the constructed libraries were sequenced using Illumina HiseqXten (Sinotech Genomics Co., Ltd.).

Revised version: After retrobeads injection, following brain slice preparation and recording, the recorded and retrobeads-labelled neurons were aspirated into the patch pipette and were then broken into the PCR tube containing 1 μl lysis buffer. For mRNA in individual cells, mRNA was amplified and cDNA by SMARTer Ultra Low Input RNA for Illumina Kit, which was qualified and reversely transcribed to cDNA by Qubit and Agilent Bioanalyzer 2100 electrophoresis. After fragmentation of cDNA (300 bp) by ultrasound, sequencing libraries (end repair, addition of poly(A), and ligation of sequencing connectors) were built using the Ovation Ultralow Library System V2. After that, the constructed libraries were sequenced using Illumina HiseqXten.

Depression models

Chronic Mild Unpredictable Stress (CMUS) procedure

The CMUS procedures consisted of food and water deprivation (24 h), day/night inversion, damp bedding (12 h), cage tilt (12 h), no bedding (12 h), rat bedding (12 h), 4°C cold bath (5 min), restraint (2 h) and tail pinching (30 min). Mice were subjected to consecutive 35 days of CMUS with two stressors per day. Non-stressed controls were handled only for cage changes and behavioral tests.

Chronic Restraint Stress (CRS) procedure

The mice were immobilized in a special restraint device for 6 h per day from 9:00 to 15:00 for 21 days.

Behavior tests

The behavior tests were performed in the following order.

Sucrose preference test

Earlier version: Mice were single-housed and trained to drink from two drinking bottles with 1 bottle of tap water and 1 bottle of 1% sucrose water, for 48 hours, and the positions of the drinking bottles were exchanged every 12 h to exclude the interference of position preference. After the training, the mice were deprived of food and water for 12 h, and then the sucrose preference test was performed for 24 h. One bottle of tap water and one bottle of 1.0 % sucrose water were given again, and the positions of the drinking bottles were exchanged at 12 h. Finally, the consumption of tap water and sugar water at 24 h were recorded to calculate the sucrose solution preference rate. Sucrose preference rate = sucrose solution consumption/(tap water consumption + sucrose solution consumption).

Revised version: Mice were single-housed and trained to drink from two drinking bottles with 1 bottle of tap water and 1 bottle of 1.0 % sucrose water, for 48 hours, and the positions of the drinking bottles were exchanged every 12 h to exclude the interference of position preference. After the training, the mice were deprived of food and water for 12 h, and then the sucrose preference test was performed for 24 h. One bottle of tap water and one bottle of 1.0 % sucrose water were given again, and the positions of the drinking bottles were exchanged at 12 h. Finally, the consumption of tap water and sugar water at 24 h were recorded to calculate the sucrose solution preference rate. Sucrose preference rate (%) = sucrose solution consumption/(tap water consumption + sucrose solution consumption)*100%.

Open field test, OFT

The mice were placed in a topless chamber (40 cm × 40 cm × 30 cm) with a camera mounted on top of the chamber and connected to ANY-maze video tracking system on the computer to automatically track and record the activity of mice during the experiment and to obtain behavioral data. The bottom of the chamber was divided into 16 small square grids (10 cm × 10 cm) on the computer, and four square grids in the central area were defined as the central zone. During the 10 min test session, total distance traveled and time spent in the center were recorded.

Elevated plus maze, EPM

Earlier version: The EPM apparatus consists of two closed arms (25 x 5 cm) across from each other and perpendicular to two open arms (25 x 5 cm) that are connected by a center platform (5 x 5 cm), with a camera in the center ceiling to automatically track and record the activity of mice. Mice were placed in the center platform facing a closed arm and allowed to freely explore the maze for 7 min, of which the first 2 min was the adaptation time. The time spent in open arms was analyzed for last 5 min after the adaptation.

Revised version: The EPM apparatus consists of two closed arms (25 × 5 cm) across from each other and perpendicular to two open arms (25 × 5 cm) that are connected by a center platform (5 × 5 cm), with a camera in the center ceiling to automatically track and record the activity of mice. Mice were placed in the center platform facing a closed arm and allowed to freely explore the maze for 7 min, of which the first 2 min was the adaptation time. The time spent in open and closed arms was analyzed for last 5 min after the adaptation.

Accelerating rotarod test

Earlier version: Locomotor activity was estimated by accelerating the rotarod test. Each group of animals was placed on an accelerating rotarod walking for 300 s from 4 to 40 rpm. After three days of training (three times per day), latency was automatically recorded from the beginning of the trial until the mouse falls off.

Revised version: Locomotor activity was estimated by accelerating the rotarod test. Each group of animals was placed on an accelerating rotarod walking for 300 s from 4 to 40 rpm. After three days of training (three times per day), latency was recorded from the beginning of the trial until the mouse falls off for 300 s from 4 to 40 rpm.

Forced swimming test, FST

Earlier version: Forced swimming test (FST) was performed in a glass cylinders, filled with the water depth about 15 cm at room temperature. The experiment was conducted for 5 min, of which the first 1 min was the adaptation time and the mice were allowed to move freely, and the immobility time was recorded for last 4 min after the adaptation.

Revised version: Forced swimming test (FST) was performed in a glass cylinder, filled with the water depth about 15 cm at room temperature. The experiment was conducted for 5 min, of which the first 1 min was the adaptation time and the mice were allowed to move freely, and the immobility time was recorded for last 4 min after the adaptation.

Tail suspension test, TST

Mouse was suspended by taping the tail (1 cm from the tip of the tail) to the fixation device for 6 min. The cumulative immobility time within last 4 min was recorded, and the time when all limbs were immobile except for respiration was considered as immobility time.

Western Blot

Earlier version: The VTA protein was isolated by100 µl RIPA and 1 µl PMSF. The total protein for each sample was transferred onto the polyvinylidene difluoride (PVDF) membranes after SDS-polyacrylamide gel electrophoresis (SDS-PAGE), and blocked 2 h at room temperature with 5% bovine serum albumin (KeyGen Biotechnology). Subsequently, all these membranes were incubated overnight at 4 °C with the primary antibodies as below: TRPC6, NALCN, and GAPDH. Following 2 h secondary antibodies incubation, all the bands were detected. The relative expression was calculated based on the internal control GAPDH.

Revised version: VTA was carefully obtained from the coronal midbrain slices (900 μm-thick) with a small homemade biopsy punch (1.0 mm in diameter) on ice, which was verified using the Franklin and Paxinos mouse brain atlas[77], and then stored at −80 °C until the day of the experiment. The VTA protein was isolated by100 µl RIPA and 1 µl PMSF. The total protein for each sample was transferred onto the polyvinylidene difluoride (PVDF) membranes after SDS-polyacrylamide gel electrophoresis (SDS-PAGE), and blocked 2 h at room temperature with 5% bovine serum albumin (KeyGen Biotechnology). Subsequently, all these membranes were incubated overnight at 4 °C with the primary antibodies as below: TRPC6, NALCN, and GAPDH. Following 2 h secondary antibodies incubation, all the bands were detected. The relative expression was calculated based on the internal control GAPDH.

Drugs and reagents

Earlier version: Drugs were bath applied at the following concentrations: 6-cyano-7-nitroquinoxaline-2,3-dione (CNQX; 10 μM; Sigma), DL-2-amino-5-phosphonopentanoic acid (APV; 50 μM; Sigma) and gabazine (10 μM; Sigma), CsCl (3 mM; sigma), ZD7288 (60 μm; Abcam), TTX (1 μM; MCE), GdCl3 (100 μM; Sigma), L703,606 (10 μM; Sigma), 2-Aminoethyldiphenylborinate (2-APB, 100 μM; Sigma), Larixyl acetate (10 μM; Sigma), Ruthenium Red (60 μM; TCI). Commercial antibodies used were: Anti-Tyrosine Hydroxylase Antibody (1:400, Millipore, MAB318, RRID: AB_2201528, for Immunofluorescence), Anti-Dopamine Transporter (N-terminal) antibody (1:400, sigma, D6944, RRID: AB_1840807, for Immunofluorescence), Anti-NALCN (1:100, Thermo Fisher Scientific, MA5-27593, RRID: AB_2735285, for Immunofluorescence), Anti-TRPC6 (1:100, Alomone, ACC-017, RRID: AB_2040243, for Immunofluorescence), Anti-GAPDH (1:10000, Santa Cruz, sc-137179,RRID: AB_2232048, for Western Blot), Anti-NALCN (1:100, GeneTex, GTX54808, for Western Blot), Anti-TRPC6 (1:100, Cell Signaling Technology, 16716, RRID:AB_2798768, for Western Blot).

Revised version: Drugs were bath applied at the following concentrations:6-cyano-7-nitroquinoxaline-2,3-dione (CNQX; 10 μM; Sigma), DL-2-amino-5-phosphonopentanoic acid (APV; 50 μM; Sigma) and gabazine (10 μM; Sigma), CsCl (3mM; sigma), ZD7288 (60 μm; Abcam), TTX (1 μM; MCE and CHENGDU MUST BIO-TECHNOLOGY CO., LTD), GdCl3 (100 μM; Sigma), L703,606 (10 μM; Sigma), 2-Aminoethyldiphenylborinate (2-APB, 100 μM; Sigma), Larixyl acetate (10 μM; Sigma), Ruthenium Red (60 μM; TCI).

Added in the revised version: Commercial antibodies used were: Anti-Tyrosine Hydroxylase Antibody (1:400, Millipore, MAB318, RRID: AB_2201528, for Immunofluorescence), Anti-Dopamine Transporter (N-terminal) antibody (1:400, sigma, D6944, RRID: AB_1840807, for Immunofluorescence), Anti-NALCN (1:100, Thermo Fisher Scientific, MA5-27593, RRID: AB_2735285, for Immunofluorescence), Anti-TRPC6 (1:100, Alomone, ACC-017, RRID: AB_2040243, for Immunofluorescence), Anti-GAPDH (1:10000, Santa Cruz, sc-137179, RRID: AB_2232048, for Western Blot), Anti-NALCN (1:100, GeneTex, GTX54808, RRID: AB_3097701, for Western Blot), Anti-TRPC6 (1:100, Cell Signaling Technology, 16716, RRID: AB_2798768, for Western Blot).

Secondary antibodies: Donkey anti-Mouse IgG (H+L) Highly Cross-Adsorbed Secondary Antibody (Alexa Fluor 488, Thermo Fisher Scientific, A-21202, 1:1000), Donkey anti-Mouse IgG (H+L) Highly Cross-Adsorbed Secondary Antibody (Alexa Fluor 546, Thermo Fisher Scientific, A10037, 1:1000), Donkey anti-Rabbit IgG (H+L) Highly Cross-Adsorbed Secondary Antibody (Alexa Fluor 488, Thermo Fisher Scientific, A21206, 1:1000), Donkey anti-Rabbit IgG (H+L) Highly Cross-Adsorbed Secondary Antibody (Alexa Fluor 546, Thermo Fisher Scientific, A10040, 1:1000), Donkey anti-Rabbit IgG (H+L) Highly Cross-Adsorbed Secondary Antibody (Alexa Fluor 647, Thermo Fisher Scientific, A31573, 1:1000)

Quantification and statistical analysis

Earlier version: Software such as GraphPad Prism6, OriginPro 8.0 (Origin Lab) and Adobe Illustrator CS6 were used for data analysis and image processing. All experimental data were expressed as mean ± standard error (mean ± S.E.M.). When the data were normally distributed, the difference between two groups was statistically analyzed by two-sample T-test or paired-sample T-test; when the data were not normally distributed, the difference between two groups was statistically analyzed by Mann-Whitney U test or Wilcoxon matched-pairs signed rank test. P < 0.05 was considered as a statistically significant difference between two groups. For the data between multiple groups, when the data were normally distributed and there was no significant variance in homogeneity, one-way ANOVA with Dunnett’s multiple comparisons test or two-way ANOVA with Sidak’s multiple comparisons test was used; when the data were not normally distributed, Kruslal-Wallis-H test with Dunnett’s multiple comparisons test and Student-Newman-Keuls test with Dunnett’s multiple comparisons test were used.

Revised version: Software such as GraphPad Prism6, OriginPro 8.0 (Origin Lab) and Adobe Illustrator CS6 were used for data analysis and image processing. All experimental data were expressed as mean ± standard error (mean ± S.E.M.). When the data were normally distributed, the difference between two groups was statistically analyzed by two-sample T-test or paired-sample T-test; when the data were not normally distributed, the difference between two groups was statistically analyzed by Mann-Whitney U test or Wilcoxon matched-pairs signed rank test. P< 0.05 was considered as a statistically significant difference between two groups. For the data between multiple groups, when the data were normally distributed and there was no significant variance inhomogeneity, one-way ANOVA with Dunnett’s multiple comparisons test or two-way ANOVA with Sidak’s multiple comparisons test was used; when the data were not normally distributed, Kruslal-Wallis-H test with Dunnett’s multiple comparisons test was used.

Removed in the revised version: Conflict of Interest

Removed in the revised version: The authors declare that the research was conducted in the absence of any commercial or financial relationships that could be construed as a potential conflict of interest.

Removed in the revised version: Author Contributions

Removed in the revised version: Hailin Zhang conceived, designed and supervised the experiments; Jing Wang performed the experiments, acquired and analyzed the data and prepared the figures; Dongmei Zhang and Min Su performed immunofluorescence of the brain slices and performed some preliminary electrophysiological experiments; Yuqi Sang and Chaoyi Li performed part of single-cell PCR; Yongxue Zhang performed part of statistical analysis; Ludi Zhang and Chenxu Niu took care of mice; Hailin Zhang and Jing Wang wrote the manuscript; Hailin Zhang and Xiaona Du prepared the final version of the manuscript.

Removed in the revised version: Funding

Removed in the revised version: This work was supported by the National Natural Science Foundation of China (81871075, 82071533) grants to HZ; National Natural Science Foundation of China (81870872) grants to XD; Science Fund for Creative Research Groups of Natural Science Foundation of Hebei Province (no. H2020206474);Basic Research Fund for Provincial Universities (JCYJ2021010), Natural Science Foundation of Hebei Province (H2023423065) and Scientific research project of Hebei administration of traditional Chinese Medicine (2024091) grants to WJ.

Supporting information